%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
Problem: data/problems/all/SYN075+1.tptp

Goal:
(?Z W. !X Y. big_f X Y <=> X = Z /\ Y = W) ==>
?W. !Y. (?Z. !X. big_f X Y <=> X = Z) <=> Y = W

Clauses:
(~big_f $X $Y \/ $X = skolemFOFtoCNF_Z) /\
(~big_f $X $Y \/ $Y = skolemFOFtoCNF_W) /\
(~($X = skolemFOFtoCNF_Z) \/ ~($Y = skolemFOFtoCNF_W) \/ big_f $X $Y) /\
(~($X = skolemFOFtoCNF_Z_1 $W) \/ skolemFOFtoCNF_Y $W = $W \/
 big_f $X (skolemFOFtoCNF_Y $W)) /\
(~(skolemFOFtoCNF_X $W $Z = $Z) \/ ~(skolemFOFtoCNF_Y $W = $W) \/
 ~big_f (skolemFOFtoCNF_X $W $Z) (skolemFOFtoCNF_Y $W)) /\
(~(skolemFOFtoCNF_Y $W = $W) \/ skolemFOFtoCNF_X $W $Z = $Z \/
 big_f (skolemFOFtoCNF_X $W $Z) (skolemFOFtoCNF_Y $W)) /\
(~big_f $X (skolemFOFtoCNF_Y $W) \/ $X = skolemFOFtoCNF_Z_1 $W \/
 skolemFOFtoCNF_Y $W = $W)

% SZS status Theorem for data/problems/all/SYN075+1.tptp

% SZS output start CNFRefutation for data/problems/all/SYN075+1.tptp
fof(pel52_1, axiom,
    (? [Z, W] : ! [X, Y] : (big_f(X, Y) <=> (X = Z & Y = W)))).

fof(pel52, conjecture,
    (? [W] : ! [Y] : (? [Z] : ! [X] : (big_f(X, Y) <=> X = Z) <=> Y = W))).

fof(subgoal_0, plain,
    (? [W] : ! [Y] : (? [Z] : ! [X] : (big_f(X, Y) <=> X = Z) <=> Y = W)),
    inference(strip, [], [pel52])).

fof(negate_0_0, plain,
    (~
       ? [W] :
         ! [Y] : (? [Z] : ! [X] : (big_f(X, Y) <=> X = Z) <=> Y = W)),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain,
    (! [W] :
       ? [Y] : (Y != W <=> ? [Z] : ! [X] : (X != Z <=> ~ big_f(X, Y)))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_1, plain,
    (! [W] :
       ? [Y] : (Y != W <=> ? [Z] : ! [X] : (X != Z <=> ~ big_f(X, Y)))),
    inference(specialize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (! [W] :
       (skolemFOFtoCNF_Y(W) != W <=>
        ? [Z] : ! [X] : (X != Z <=> ~ big_f(X, skolemFOFtoCNF_Y(W))))),
    inference(skolemize, [], [normalize_0_1])).

fof(normalize_0_3, plain,
    (! [W, X, Z] :
       ((X != skolemFOFtoCNF_Z_1(W) | skolemFOFtoCNF_Y(W) = W |
         big_f(X, skolemFOFtoCNF_Y(W))) &
        (skolemFOFtoCNF_X(W, Z) != Z | skolemFOFtoCNF_Y(W) != W |
         ~ big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))) &
        (skolemFOFtoCNF_Y(W) != W | skolemFOFtoCNF_X(W, Z) = Z |
         big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))) &
        (~ big_f(X, skolemFOFtoCNF_Y(W)) | X = skolemFOFtoCNF_Z_1(W) |
         skolemFOFtoCNF_Y(W) = W))),
    inference(clausify, [], [normalize_0_2])).

fof(normalize_0_4, plain,
    (! [W, Z] :
       (skolemFOFtoCNF_X(W, Z) != Z | skolemFOFtoCNF_Y(W) != W |
        ~ big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W)))),
    inference(conjunct, [], [normalize_0_3])).

fof(normalize_0_5, plain,
    (? [W, Z] : ! [X, Y] : (~ big_f(X, Y) <=> (X != Z | Y != W))),
    inference(canonicalize, [], [pel52_1])).

fof(normalize_0_6, plain,
    (! [X, Y] :
       (~ big_f(X, Y) <=>
        (X != skolemFOFtoCNF_Z | Y != skolemFOFtoCNF_W))),
    inference(skolemize, [], [normalize_0_5])).

fof(normalize_0_7, plain,
    (! [X, Y] :
       (~ big_f(X, Y) <=>
        (X != skolemFOFtoCNF_Z | Y != skolemFOFtoCNF_W))),
    inference(specialize, [], [normalize_0_6])).

fof(normalize_0_8, plain,
    (! [X, Y] :
       ((~ big_f(X, Y) | X = skolemFOFtoCNF_Z) &
        (~ big_f(X, Y) | Y = skolemFOFtoCNF_W) &
        (X != skolemFOFtoCNF_Z | Y != skolemFOFtoCNF_W | big_f(X, Y)))),
    inference(clausify, [], [normalize_0_7])).

fof(normalize_0_9, plain,
    (! [X, Y] : (~ big_f(X, Y) | X = skolemFOFtoCNF_Z)),
    inference(conjunct, [], [normalize_0_8])).

fof(normalize_0_10, plain,
    (! [W, Z] :
       (skolemFOFtoCNF_Y(W) != W | skolemFOFtoCNF_X(W, Z) = Z |
        big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W)))),
    inference(conjunct, [], [normalize_0_3])).

fof(normalize_0_11, plain,
    (! [W, X] :
       (X != skolemFOFtoCNF_Z_1(W) | skolemFOFtoCNF_Y(W) = W |
        big_f(X, skolemFOFtoCNF_Y(W)))),
    inference(conjunct, [], [normalize_0_3])).

fof(normalize_0_12, plain,
    (! [X, Y] : (~ big_f(X, Y) | Y = skolemFOFtoCNF_W)),
    inference(conjunct, [], [normalize_0_8])).

fof(normalize_0_13, plain,
    (! [X, Y] :
       (X != skolemFOFtoCNF_Z | Y != skolemFOFtoCNF_W | big_f(X, Y))),
    inference(conjunct, [], [normalize_0_8])).

cnf(refute_0_0, plain,
    (skolemFOFtoCNF_X(W, Z) != Z | skolemFOFtoCNF_Y(W) != W |
     ~ big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))),
    inference(canonicalize, [], [normalize_0_4])).

cnf(refute_0_1, plain,
    (skolemFOFtoCNF_X(W, Z) != Z | ~ big_f(Z, skolemFOFtoCNF_Y(W)) |
     big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))),
    introduced(tautology,
               [equality,
                [$cnf(~ big_f(skolemFOFtoCNF_X(W, Z),
                        skolemFOFtoCNF_Y(W))), [0], $fot(Z)]])).

cnf(refute_0_2, plain,
    (skolemFOFtoCNF_Y(W) != W | ~ big_f(Z, W) |
     big_f(Z, skolemFOFtoCNF_Y(W))),
    introduced(tautology,
               [equality,
                [$cnf(~ big_f(Z, skolemFOFtoCNF_Y(W))), [1], $fot(W)]])).

cnf(refute_0_3, plain,
    (skolemFOFtoCNF_X(W, Z) != Z | skolemFOFtoCNF_Y(W) != W |
     ~ big_f(Z, W) | big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))),
    inference(resolve, [$cnf(big_f(Z, skolemFOFtoCNF_Y(W)))],
              [refute_0_2, refute_0_1])).

cnf(refute_0_4, plain,
    (skolemFOFtoCNF_X(W, Z) != Z | skolemFOFtoCNF_Y(W) != W |
     ~ big_f(Z, W)),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W)))],
              [refute_0_3, refute_0_0])).

cnf(refute_0_5, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, skolemFOFtoCNF_Z) !=
     skolemFOFtoCNF_Z |
     skolemFOFtoCNF_Y(skolemFOFtoCNF_W) != skolemFOFtoCNF_W |
     ~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(subst, [],
              [refute_0_4 :
               [bind(W, $fot(skolemFOFtoCNF_W)),
                bind(Z, $fot(skolemFOFtoCNF_Z))]])).

cnf(refute_0_6, plain, (~ big_f(X, Y) | X = skolemFOFtoCNF_Z),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_7, plain,
    (~ big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21), skolemFOFtoCNF_W) |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21) = skolemFOFtoCNF_Z),
    inference(subst, [],
              [refute_0_6 :
               [bind(X, $fot(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21))),
                bind(Y, $fot(skolemFOFtoCNF_W))]])).

cnf(refute_0_8, plain,
    (skolemFOFtoCNF_Y(W) != W | skolemFOFtoCNF_X(W, Z) = Z |
     big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))),
    inference(canonicalize, [], [normalize_0_10])).

cnf(refute_0_9, plain,
    (skolemFOFtoCNF_Y(W) != W |
     ~ big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W)) |
     big_f(skolemFOFtoCNF_X(W, Z), W)),
    introduced(tautology,
               [equality,
                [$cnf(big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W))),
                 [1], $fot(W)]])).

cnf(refute_0_10, plain,
    (skolemFOFtoCNF_Y(W) != W | skolemFOFtoCNF_X(W, Z) = Z |
     big_f(skolemFOFtoCNF_X(W, Z), W)),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_X(W, Z), skolemFOFtoCNF_Y(W)))],
              [refute_0_8, refute_0_9])).

cnf(refute_0_11, plain,
    (skolemFOFtoCNF_Y(skolemFOFtoCNF_W) != skolemFOFtoCNF_W |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20) = X_20 |
     big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20), skolemFOFtoCNF_W)),
    inference(subst, [],
              [refute_0_10 :
               [bind(W, $fot(skolemFOFtoCNF_W)), bind(Z, $fot(X_20))]])).

cnf(refute_0_12, plain,
    (X != skolemFOFtoCNF_Z_1(W) | skolemFOFtoCNF_Y(W) = W |
     big_f(X, skolemFOFtoCNF_Y(W))),
    inference(canonicalize, [], [normalize_0_11])).

cnf(refute_0_13, plain,
    (skolemFOFtoCNF_Z_1(W) != skolemFOFtoCNF_Z_1(W) |
     skolemFOFtoCNF_Y(W) = W |
     big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W))),
    inference(subst, [],
              [refute_0_12 : [bind(X, $fot(skolemFOFtoCNF_Z_1(W)))]])).

cnf(refute_0_14, plain, (skolemFOFtoCNF_Z_1(W) = skolemFOFtoCNF_Z_1(W)),
    introduced(tautology, [refl, [$fot(skolemFOFtoCNF_Z_1(W))]])).

cnf(refute_0_15, plain,
    (skolemFOFtoCNF_Y(W) = W |
     big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W))),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Z_1(W)))],
              [refute_0_14, refute_0_13])).

cnf(refute_0_16, plain,
    (~ big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W)) |
     skolemFOFtoCNF_Z_1(W) = skolemFOFtoCNF_Z),
    inference(subst, [],
              [refute_0_6 :
               [bind(X, $fot(skolemFOFtoCNF_Z_1(W))),
                bind(Y, $fot(skolemFOFtoCNF_Y(W)))]])).

cnf(refute_0_17, plain,
    (skolemFOFtoCNF_Y(W) = W | skolemFOFtoCNF_Z_1(W) = skolemFOFtoCNF_Z),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W)))],
              [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (skolemFOFtoCNF_Z_1(W) != skolemFOFtoCNF_Z |
     ~ big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W)) |
     big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_Y(W))),
    introduced(tautology,
               [equality,
                [$cnf(big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W))),
                 [0], $fot(skolemFOFtoCNF_Z)]])).

cnf(refute_0_19, plain,
    (~ big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W)) |
     skolemFOFtoCNF_Y(W) = W |
     big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_Y(W))),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Z))],
              [refute_0_17, refute_0_18])).

cnf(refute_0_20, plain,
    (skolemFOFtoCNF_Y(W) = W |
     big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_Y(W))),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_Z_1(W), skolemFOFtoCNF_Y(W)))],
              [refute_0_15, refute_0_19])).

cnf(refute_0_21, plain, (~ big_f(X, Y) | Y = skolemFOFtoCNF_W),
    inference(canonicalize, [], [normalize_0_12])).

cnf(refute_0_22, plain,
    (~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_Y(W)) |
     skolemFOFtoCNF_Y(W) = skolemFOFtoCNF_W),
    inference(subst, [],
              [refute_0_21 :
               [bind(X, $fot(skolemFOFtoCNF_Z)),
                bind(Y, $fot(skolemFOFtoCNF_Y(W)))]])).

cnf(refute_0_23, plain,
    (skolemFOFtoCNF_Y(W) = W | skolemFOFtoCNF_Y(W) = skolemFOFtoCNF_W),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_Y(W)))],
              [refute_0_20, refute_0_22])).

cnf(refute_0_24, plain,
    (skolemFOFtoCNF_Y(skolemFOFtoCNF_W) = skolemFOFtoCNF_W),
    inference(subst, [],
              [refute_0_23 : [bind(W, $fot(skolemFOFtoCNF_W))]])).

cnf(refute_0_25, plain,
    (skolemFOFtoCNF_W != skolemFOFtoCNF_W |
     skolemFOFtoCNF_Y(skolemFOFtoCNF_W) != skolemFOFtoCNF_W |
     skolemFOFtoCNF_Y(skolemFOFtoCNF_W) = skolemFOFtoCNF_W),
    introduced(tautology,
               [equality,
                [$cnf($equal(skolemFOFtoCNF_Y(skolemFOFtoCNF_W),
                        skolemFOFtoCNF_W)), [0, 0],
                 $fot(skolemFOFtoCNF_W)]])).

cnf(refute_0_26, plain,
    (skolemFOFtoCNF_W != skolemFOFtoCNF_W |
     skolemFOFtoCNF_Y(skolemFOFtoCNF_W) = skolemFOFtoCNF_W),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Y(skolemFOFtoCNF_W),
                      skolemFOFtoCNF_W))], [refute_0_24, refute_0_25])).

cnf(refute_0_27, plain,
    (skolemFOFtoCNF_W != skolemFOFtoCNF_W |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20) = X_20 |
     big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20), skolemFOFtoCNF_W)),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Y(skolemFOFtoCNF_W),
                      skolemFOFtoCNF_W))], [refute_0_26, refute_0_11])).

cnf(refute_0_28, plain, (skolemFOFtoCNF_W = skolemFOFtoCNF_W),
    introduced(tautology, [refl, [$fot(skolemFOFtoCNF_W)]])).

cnf(refute_0_29, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20) = X_20 |
     big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_20), skolemFOFtoCNF_W)),
    inference(resolve, [$cnf($equal(skolemFOFtoCNF_W, skolemFOFtoCNF_W))],
              [refute_0_28, refute_0_27])).

cnf(refute_0_30, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21) = X_21 |
     big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21), skolemFOFtoCNF_W)),
    inference(subst, [], [refute_0_29 : [bind(X_20, $fot(X_21))]])).

cnf(refute_0_31, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21) = X_21 |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21) = skolemFOFtoCNF_Z),
    inference(resolve,
              [$cnf(big_f(skolemFOFtoCNF_X(skolemFOFtoCNF_W, X_21),
                      skolemFOFtoCNF_W))], [refute_0_30, refute_0_7])).

cnf(refute_0_32, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, skolemFOFtoCNF_Z) =
     skolemFOFtoCNF_Z),
    inference(subst, [],
              [refute_0_31 : [bind(X_21, $fot(skolemFOFtoCNF_Z))]])).

cnf(refute_0_33, plain,
    (skolemFOFtoCNF_X(skolemFOFtoCNF_W, skolemFOFtoCNF_Z) !=
     skolemFOFtoCNF_Z | skolemFOFtoCNF_Z != skolemFOFtoCNF_Z |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, skolemFOFtoCNF_Z) =
     skolemFOFtoCNF_Z),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(skolemFOFtoCNF_X(skolemFOFtoCNF_W,
                          skolemFOFtoCNF_Z), skolemFOFtoCNF_Z)), [0],
                 $fot(skolemFOFtoCNF_Z)]])).

cnf(refute_0_34, plain,
    (skolemFOFtoCNF_Z != skolemFOFtoCNF_Z |
     skolemFOFtoCNF_X(skolemFOFtoCNF_W, skolemFOFtoCNF_Z) =
     skolemFOFtoCNF_Z),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_X(skolemFOFtoCNF_W,
                        skolemFOFtoCNF_Z), skolemFOFtoCNF_Z))],
              [refute_0_32, refute_0_33])).

cnf(refute_0_35, plain,
    (skolemFOFtoCNF_Y(skolemFOFtoCNF_W) != skolemFOFtoCNF_W |
     skolemFOFtoCNF_Z != skolemFOFtoCNF_Z |
     ~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_X(skolemFOFtoCNF_W,
                        skolemFOFtoCNF_Z), skolemFOFtoCNF_Z))],
              [refute_0_34, refute_0_5])).

cnf(refute_0_36, plain, (skolemFOFtoCNF_Z = skolemFOFtoCNF_Z),
    introduced(tautology, [refl, [$fot(skolemFOFtoCNF_Z)]])).

cnf(refute_0_37, plain,
    (skolemFOFtoCNF_Y(skolemFOFtoCNF_W) != skolemFOFtoCNF_W |
     ~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve, [$cnf($equal(skolemFOFtoCNF_Z, skolemFOFtoCNF_Z))],
              [refute_0_36, refute_0_35])).

cnf(refute_0_38, plain,
    (skolemFOFtoCNF_W != skolemFOFtoCNF_W |
     ~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_Y(skolemFOFtoCNF_W),
                      skolemFOFtoCNF_W))], [refute_0_26, refute_0_37])).

cnf(refute_0_39, plain, (~ big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve, [$cnf($equal(skolemFOFtoCNF_W, skolemFOFtoCNF_W))],
              [refute_0_28, refute_0_38])).

cnf(refute_0_40, plain,
    (X != skolemFOFtoCNF_Z | Y != skolemFOFtoCNF_W | big_f(X, Y)),
    inference(canonicalize, [], [normalize_0_13])).

cnf(refute_0_41, plain,
    (skolemFOFtoCNF_W != skolemFOFtoCNF_W |
     skolemFOFtoCNF_Z != skolemFOFtoCNF_Z |
     big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(subst, [],
              [refute_0_40 :
               [bind(X, $fot(skolemFOFtoCNF_Z)),
                bind(Y, $fot(skolemFOFtoCNF_W))]])).

cnf(refute_0_42, plain,
    (skolemFOFtoCNF_Z != skolemFOFtoCNF_Z |
     big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve, [$cnf($equal(skolemFOFtoCNF_W, skolemFOFtoCNF_W))],
              [refute_0_28, refute_0_41])).

cnf(refute_0_43, plain, (big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W)),
    inference(resolve, [$cnf($equal(skolemFOFtoCNF_Z, skolemFOFtoCNF_Z))],
              [refute_0_36, refute_0_42])).

cnf(refute_0_44, plain, ($false),
    inference(resolve, [$cnf(big_f(skolemFOFtoCNF_Z, skolemFOFtoCNF_W))],
              [refute_0_43, refute_0_39])).
% SZS output end CNFRefutation for data/problems/all/SYN075+1.tptp

%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%%
Problem: data/problems/all/SEU140+2.tptp

Goal:
(!A B. in A B ==> ~in B A) /\
(!A B. proper_subset A B ==> ~proper_subset B A) /\
(!A B. set_union2 A B = set_union2 B A) /\
(!A B. set_intersection2 A B = set_intersection2 B A) /\
(!A B. A = B <=> subset A B /\ subset B A) /\
(!A. A = empty_set <=> !B. ~in B A) /\
(!A B C. C = set_union2 A B <=> !D. in D C <=> in D A \/ in D B) /\
(!A B. subset A B <=> !C. in C A ==> in C B) /\
(!A B C. C = set_intersection2 A B <=> !D. in D C <=> in D A /\ in D B) /\
(!A B C. C = set_difference A B <=> !D. in D C <=> in D A /\ ~in D B) /\
(!A B. disjoint A B <=> set_intersection2 A B = empty_set) /\
(!A B. proper_subset A B <=> subset A B /\ ~(A = B)) /\ T /\ T /\ T /\ T /\
empty empty_set /\ (!A B. ~empty A ==> ~empty (set_union2 A B)) /\
(!A B. ~empty A ==> ~empty (set_union2 B A)) /\
(!A B. set_union2 A A = A) /\ (!A B. set_intersection2 A A = A) /\
(!A B. ~proper_subset A A) /\
(!A B. set_difference A B = empty_set <=> subset A B) /\ (?A. empty A) /\
(?A. ~empty A) /\ (!A B. subset A A) /\
(!A B. disjoint A B ==> disjoint B A) /\
(!A B. subset A B ==> set_union2 A B = B) /\
(!A B. subset (set_intersection2 A B) A) /\
(!A B C. subset A B /\ subset A C ==> subset A (set_intersection2 B C)) /\
(!A. set_union2 A empty_set = A) /\
(!A B C. subset A B /\ subset B C ==> subset A C) /\
(!A B C.
   subset A B ==>
   subset (set_intersection2 A C) (set_intersection2 B C)) /\
(!A B. subset A B ==> set_intersection2 A B = A) /\
(!A. set_intersection2 A empty_set = empty_set) /\
(!A B. (!C. in C A <=> in C B) ==> A = B) /\ (!A. subset empty_set A) /\
(!A B C.
   subset A B ==> subset (set_difference A C) (set_difference B C)) /\
(!A B. subset (set_difference A B) A) /\
(!A B. set_difference A B = empty_set <=> subset A B) /\
(!A B. set_union2 A (set_difference B A) = set_union2 A B) /\
(!A. set_difference A empty_set = A) /\
(!A B.
   ~(~disjoint A B /\ !C. ~(in C A /\ in C B)) /\
   ~((?C. in C A /\ in C B) /\ disjoint A B)) /\
(!A. subset A empty_set ==> A = empty_set) /\
(!A B. set_difference (set_union2 A B) B = set_difference A B) /\
(!A B. subset A B ==> B = set_union2 A (set_difference B A)) /\
(!A B. set_difference A (set_difference A B) = set_intersection2 A B) /\
(!A. set_difference empty_set A = empty_set) /\
(!A B.
   ~(~disjoint A B /\ !C. ~in C (set_intersection2 A B)) /\
   ~((?C. in C (set_intersection2 A B)) /\ disjoint A B)) /\
(!A B. ~(subset A B /\ proper_subset B A)) /\
(!A. empty A ==> A = empty_set) /\ (!A B. ~(in A B /\ empty B)) /\
(!A B. subset A (set_union2 A B)) /\
(!A B. ~(empty A /\ ~(A = B) /\ empty B)) /\
(!A B C. subset A B /\ subset C B ==> subset (set_union2 A C) B) ==>
!A B C. subset A B /\ disjoint B C ==> disjoint A C

Clauses:
(~in $A $B \/ ~in $B $A) /\
(~proper_subset $A $B \/ ~proper_subset $B $A) /\
set_union2 $A $B = set_union2 $B $A /\
set_intersection2 $A $B = set_intersection2 $B $A /\
(~($A = $B) \/ subset $A $B) /\ (~($A = $B) \/ subset $B $A) /\
(~subset $A $B \/ ~subset $B $A \/ $A = $B) /\
(~($A = empty_set) \/ ~in $B $A) /\
($A = empty_set \/ in (skolemFOFtoCNF_B $A) $A) /\
(~($C = set_union2 $A $B) \/ ~in $D $A \/ in $D $C) /\
(~($C = set_union2 $A $B) \/ ~in $D $B \/ in $D $C) /\
(~in (skolemFOFtoCNF_D $A $B $C) $A \/
 ~in (skolemFOFtoCNF_D $A $B $C) $C \/ $C = set_union2 $A $B) /\
(~in (skolemFOFtoCNF_D $A $B $C) $B \/
 ~in (skolemFOFtoCNF_D $A $B $C) $C \/ $C = set_union2 $A $B) /\
(~($C = set_union2 $A $B) \/ ~in $D $C \/ in $D $A \/ in $D $B) /\
($C = set_union2 $A $B \/ in (skolemFOFtoCNF_D $A $B $C) $A \/
 in (skolemFOFtoCNF_D $A $B $C) $B \/ in (skolemFOFtoCNF_D $A $B $C) $C) /\
(~in (skolemFOFtoCNF_C $A $B) $B \/ subset $A $B) /\
(in (skolemFOFtoCNF_C $A $B) $A \/ subset $A $B) /\
(~in $C $A \/ ~subset $A $B \/ in $C $B) /\
(~($C = set_intersection2 $A $B) \/ ~in $D $C \/ in $D $A) /\
(~($C = set_intersection2 $A $B) \/ ~in $D $C \/ in $D $B) /\
($C = set_intersection2 $A $B \/ in (skolemFOFtoCNF_D_1 $A $B $C) $A \/
 in (skolemFOFtoCNF_D_1 $A $B $C) $C) /\
($C = set_intersection2 $A $B \/ in (skolemFOFtoCNF_D_1 $A $B $C) $B \/
 in (skolemFOFtoCNF_D_1 $A $B $C) $C) /\
(~($C = set_intersection2 $A $B) \/ ~in $D $A \/ ~in $D $B \/ in $D $C) /\
(~in (skolemFOFtoCNF_D_1 $A $B $C) $A \/
 ~in (skolemFOFtoCNF_D_1 $A $B $C) $B \/
 ~in (skolemFOFtoCNF_D_1 $A $B $C) $C \/ $C = set_intersection2 $A $B) /\
(~($C = set_difference $A $B) \/ ~in $D $B \/ ~in $D $C) /\
(~($C = set_difference $A $B) \/ ~in $D $C \/ in $D $A) /\
(~in (skolemFOFtoCNF_D_2 $A $B $C) $B \/ $C = set_difference $A $B \/
 in (skolemFOFtoCNF_D_2 $A $B $C) $C) /\
($C = set_difference $A $B \/ in (skolemFOFtoCNF_D_2 $A $B $C) $A \/
 in (skolemFOFtoCNF_D_2 $A $B $C) $C) /\
(~($C = set_difference $A $B) \/ ~in $D $A \/ in $D $B \/ in $D $C) /\
(~in (skolemFOFtoCNF_D_2 $A $B $C) $A \/
 ~in (skolemFOFtoCNF_D_2 $A $B $C) $C \/ $C = set_difference $A $B \/
 in (skolemFOFtoCNF_D_2 $A $B $C) $B) /\
(~(set_intersection2 $A $B = empty_set) \/ disjoint $A $B) /\
(~disjoint $A $B \/ set_intersection2 $A $B = empty_set) /\
(~($A = $B) \/ ~proper_subset $A $B) /\
(~proper_subset $A $B \/ subset $A $B) /\
(~subset $A $B \/ $A = $B \/ proper_subset $A $B) /\ empty empty_set /\
(~empty (set_union2 $A $B) \/ empty $A) /\
(~empty (set_union2 $B $A) \/ empty $A) /\ set_union2 $A $A = $A /\
set_intersection2 $A $A = $A /\ ~proper_subset $A $A /\
(~(set_difference $A $B = empty_set) \/ subset $A $B) /\
(~subset $A $B \/ set_difference $A $B = empty_set) /\
empty skolemFOFtoCNF_A /\ ~empty skolemFOFtoCNF_A_1 /\ subset $A $A /\
(~disjoint $A $B \/ disjoint $B $A) /\
(~subset $A $B \/ set_union2 $A $B = $B) /\
subset (set_intersection2 $A $B) $A /\
(~subset $A $B \/ ~subset $A $C \/ subset $A (set_intersection2 $B $C)) /\
set_union2 $A empty_set = $A /\
(~subset $A $B \/ ~subset $B $C \/ subset $A $C) /\
(~subset $A $B \/
 subset (set_intersection2 $A $C) (set_intersection2 $B $C)) /\
(~subset $A $B \/ set_intersection2 $A $B = $A) /\
set_intersection2 $A empty_set = empty_set /\
(~in (skolemFOFtoCNF_C_1 $A $B) $A \/ ~in (skolemFOFtoCNF_C_1 $A $B) $B \/
 $A = $B) /\
($A = $B \/ in (skolemFOFtoCNF_C_1 $A $B) $A \/
 in (skolemFOFtoCNF_C_1 $A $B) $B) /\ subset empty_set $A /\
(~subset $A $B \/ subset (set_difference $A $C) (set_difference $B $C)) /\
subset (set_difference $A $B) $A /\
set_union2 $A (set_difference $B $A) = set_union2 $A $B /\
set_difference $A empty_set = $A /\
(~disjoint $A $B \/ ~in $C $A \/ ~in $C $B) /\
(disjoint $A $B \/ in (skolemFOFtoCNF_C_2 $A $B) $A) /\
(disjoint $A $B \/ in (skolemFOFtoCNF_C_2 $A $B) $B) /\
(~subset $A empty_set \/ $A = empty_set) /\
set_difference (set_union2 $A $B) $B = set_difference $A $B /\
(~subset $A $B \/ $B = set_union2 $A (set_difference $B $A)) /\
set_difference $A (set_difference $A $B) = set_intersection2 $A $B /\
set_difference empty_set $A = empty_set /\
(~disjoint $A $B \/ ~in $C (set_intersection2 $A $B)) /\
(disjoint $A $B \/
 in (skolemFOFtoCNF_C_3 $A $B) (set_intersection2 $A $B)) /\
(~proper_subset $B $A \/ ~subset $A $B) /\ (~empty $A \/ $A = empty_set) /\
(~empty $B \/ ~in $A $B) /\ subset $A (set_union2 $A $B) /\
(~empty $A \/ ~empty $B \/ $A = $B) /\
(~subset $A $B \/ ~subset $C $B \/ subset (set_union2 $A $C) $B) /\
~disjoint skolemFOFtoCNF_A_2 skolemFOFtoCNF_C_4 /\
disjoint skolemFOFtoCNF_B_1 skolemFOFtoCNF_C_4 /\
subset skolemFOFtoCNF_A_2 skolemFOFtoCNF_B_1

% SZS status Theorem for data/problems/all/SEU140+2.tptp

% SZS output start CNFRefutation for data/problems/all/SEU140+2.tptp
fof(commutativity_k2_xboole_0, axiom,
    (! [A, B] : set_union2(A, B) = set_union2(B, A))).

fof(commutativity_k3_xboole_0, axiom,
    (! [A, B] : set_intersection2(A, B) = set_intersection2(B, A))).

fof(d3_xboole_0, axiom,
    (! [A, B, C] :
       (C = set_intersection2(A, B) <=>
        ! [D] : (in(D, C) <=> (in(D, A) & in(D, B)))))).

fof(d4_xboole_0, axiom,
    (! [A, B, C] :
       (C = set_difference(A, B) <=>
        ! [D] : (in(D, C) <=> (in(D, A) & ~ in(D, B)))))).

fof(d7_xboole_0, axiom,
    (! [A, B] : (disjoint(A, B) <=> set_intersection2(A, B) = empty_set))).

fof(l32_xboole_1, lemma,
    (! [A, B] : (set_difference(A, B) = empty_set <=> subset(A, B)))).

fof(symmetry_r1_xboole_0, axiom,
    (! [A, B] : (disjoint(A, B) => disjoint(B, A)))).

fof(t1_boole, axiom, (! [A] : set_union2(A, empty_set) = A)).

fof(t36_xboole_1, lemma, (! [A, B] : subset(set_difference(A, B), A))).

fof(t37_xboole_1, lemma,
    (! [A, B] : (set_difference(A, B) = empty_set <=> subset(A, B)))).

fof(t39_xboole_1, lemma,
    (! [A, B] : set_union2(A, set_difference(B, A)) = set_union2(A, B))).

fof(t3_boole, axiom, (! [A] : set_difference(A, empty_set) = A)).

fof(t40_xboole_1, lemma,
    (! [A, B] :
       set_difference(set_union2(A, B), B) = set_difference(A, B))).

fof(t48_xboole_1, lemma,
    (! [A, B] :
       set_difference(A, set_difference(A, B)) = set_intersection2(A, B))).

fof(t4_boole, axiom, (! [A] : set_difference(empty_set, A) = empty_set)).

fof(t4_xboole_0, lemma,
    (! [A, B] :
       (~ (~ disjoint(A, B) & ! [C] : ~ in(C, set_intersection2(A, B))) &
        ~ (? [C] : in(C, set_intersection2(A, B)) & disjoint(A, B))))).

fof(t63_xboole_1, conjecture,
    (! [A, B, C] : ((subset(A, B) & disjoint(B, C)) => disjoint(A, C)))).

fof(subgoal_0, plain,
    (! [A, B, C] : ((subset(A, B) & disjoint(B, C)) => disjoint(A, C))),
    inference(strip, [], [t63_xboole_1])).

fof(negate_0_0, plain,
    (~ ! [A, B, C] : ((subset(A, B) & disjoint(B, C)) => disjoint(A, C))),
    inference(negate, [], [subgoal_0])).

fof(normalize_0_0, plain, (! [A, B] : (~ disjoint(A, B) | disjoint(B, A))),
    inference(canonicalize, [], [symmetry_r1_xboole_0])).

fof(normalize_0_1, plain, (! [A, B] : (~ disjoint(A, B) | disjoint(B, A))),
    inference(specialize, [], [normalize_0_0])).

fof(normalize_0_2, plain,
    (! [A, B, C] :
       (C != set_difference(A, B) <=>
        ? [D] : (~ in(D, C) <=> (~ in(D, B) & in(D, A))))),
    inference(canonicalize, [], [d4_xboole_0])).

fof(normalize_0_3, plain,
    (! [A, B, C] :
       (C != set_difference(A, B) <=>
        ? [D] : (~ in(D, C) <=> (~ in(D, B) & in(D, A))))),
    inference(specialize, [], [normalize_0_2])).

fof(normalize_0_4, plain,
    (! [A, B, C, D] :
       ((C != set_difference(A, B) | ~ in(D, B) | ~ in(D, C)) &
        (C != set_difference(A, B) | ~ in(D, C) | in(D, A)) &
        (~ in(skolemFOFtoCNF_D_2(A, B, C), B) | C = set_difference(A, B) |
         in(skolemFOFtoCNF_D_2(A, B, C), C)) &
        (C = set_difference(A, B) | in(skolemFOFtoCNF_D_2(A, B, C), A) |
         in(skolemFOFtoCNF_D_2(A, B, C), C)) &
        (C != set_difference(A, B) | ~ in(D, A) | in(D, B) | in(D, C)) &
        (~ in(skolemFOFtoCNF_D_2(A, B, C), A) |
         ~ in(skolemFOFtoCNF_D_2(A, B, C), C) | C = set_difference(A, B) |
         in(skolemFOFtoCNF_D_2(A, B, C), B)))),
    inference(clausify, [], [normalize_0_3])).

fof(normalize_0_5, plain,
    (! [A, B, C, D] :
       (C != set_difference(A, B) | ~ in(D, B) | ~ in(D, C))),
    inference(conjunct, [], [normalize_0_4])).

fof(normalize_0_6, plain,
    (! [A, B] : set_union2(A, set_difference(B, A)) = set_union2(A, B)),
    inference(canonicalize, [], [t39_xboole_1])).

fof(normalize_0_7, plain,
    (! [A, B] : set_union2(A, set_difference(B, A)) = set_union2(A, B)),
    inference(specialize, [], [normalize_0_6])).

fof(normalize_0_8, plain,
    (! [A, B] :
       set_difference(A, set_difference(A, B)) = set_intersection2(A, B)),
    inference(canonicalize, [], [t48_xboole_1])).

fof(normalize_0_9, plain,
    (! [A, B] :
       set_difference(A, set_difference(A, B)) = set_intersection2(A, B)),
    inference(specialize, [], [normalize_0_8])).

fof(normalize_0_10, plain,
    (! [A, B] : set_union2(A, B) = set_union2(B, A)),
    inference(canonicalize, [], [commutativity_k2_xboole_0])).

fof(normalize_0_11, plain,
    (! [A, B] : set_union2(A, B) = set_union2(B, A)),
    inference(specialize, [], [normalize_0_10])).

fof(normalize_0_12, plain,
    (? [A, B, C] : (~ disjoint(A, C) & disjoint(B, C) & subset(A, B))),
    inference(canonicalize, [], [negate_0_0])).

fof(normalize_0_13, plain,
    (~ disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4) &
     disjoint(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) &
     subset(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(skolemize, [], [normalize_0_12])).

fof(normalize_0_14, plain,
    (disjoint(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(conjunct, [], [normalize_0_13])).

fof(normalize_0_15, plain,
    (! [A, B] :
       (set_intersection2(A, B) != empty_set <=> ~ disjoint(A, B))),
    inference(canonicalize, [], [d7_xboole_0])).

fof(normalize_0_16, plain,
    (! [A, B] :
       (set_intersection2(A, B) != empty_set <=> ~ disjoint(A, B))),
    inference(specialize, [], [normalize_0_15])).

fof(normalize_0_17, plain,
    (! [A, B] :
       ((set_intersection2(A, B) != empty_set | disjoint(A, B)) &
        (~ disjoint(A, B) | set_intersection2(A, B) = empty_set))),
    inference(clausify, [], [normalize_0_16])).

fof(normalize_0_18, plain,
    (! [A, B] : (~ disjoint(A, B) | set_intersection2(A, B) = empty_set)),
    inference(conjunct, [], [normalize_0_17])).

fof(normalize_0_19, plain, (! [A] : set_union2(A, empty_set) = A),
    inference(canonicalize, [], [t1_boole])).

fof(normalize_0_20, plain, (! [A] : set_union2(A, empty_set) = A),
    inference(specialize, [], [normalize_0_19])).

fof(normalize_0_21, plain,
    (! [A, B] :
       set_difference(set_union2(A, B), B) = set_difference(A, B)),
    inference(canonicalize, [], [t40_xboole_1])).

fof(normalize_0_22, plain,
    (! [A, B] :
       set_difference(set_union2(A, B), B) = set_difference(A, B)),
    inference(specialize, [], [normalize_0_21])).

fof(normalize_0_23, plain,
    (! [A, B] : set_intersection2(A, B) = set_intersection2(B, A)),
    inference(canonicalize, [], [commutativity_k3_xboole_0])).

fof(normalize_0_24, plain,
    (! [A, B] : set_intersection2(A, B) = set_intersection2(B, A)),
    inference(specialize, [], [normalize_0_23])).

fof(normalize_0_25, plain,
    (! [A] : set_difference(empty_set, A) = empty_set),
    inference(canonicalize, [], [t4_boole])).

fof(normalize_0_26, plain,
    (! [A] : set_difference(empty_set, A) = empty_set),
    inference(specialize, [], [normalize_0_25])).

fof(normalize_0_27, plain, (! [A] : set_difference(A, empty_set) = A),
    inference(canonicalize, [], [t3_boole])).

fof(normalize_0_28, plain, (! [A] : set_difference(A, empty_set) = A),
    inference(specialize, [], [normalize_0_27])).

fof(normalize_0_29, plain, (! [A, B] : subset(set_difference(A, B), A)),
    inference(canonicalize, [], [t36_xboole_1])).

fof(normalize_0_30, plain, (! [A, B] : subset(set_difference(A, B), A)),
    inference(specialize, [], [normalize_0_29])).

fof(normalize_0_31, plain,
    (! [A, B] : (set_difference(A, B) != empty_set <=> ~ subset(A, B))),
    inference(canonicalize, [], [t37_xboole_1])).

fof(normalize_0_32, plain,
    (! [A, B] : (set_difference(A, B) != empty_set <=> ~ subset(A, B))),
    inference(specialize, [], [normalize_0_31])).

fof(normalize_0_33, plain,
    (! [A, B] :
       ((set_difference(A, B) != empty_set | subset(A, B)) &
        (~ subset(A, B) | set_difference(A, B) = empty_set))),
    inference(clausify, [], [normalize_0_32])).

fof(normalize_0_34, plain,
    (! [A, B] : (~ subset(A, B) | set_difference(A, B) = empty_set)),
    inference(conjunct, [], [normalize_0_33])).

fof(normalize_0_35, plain,
    (! [A, B] :
       (~ disjoint(A, B) | ! [C] : ~ in(C, set_intersection2(A, B))) &
     ! [A, B] : (disjoint(A, B) | ? [C] : in(C, set_intersection2(A, B)))),
    inference(canonicalize, [], [t4_xboole_0])).

fof(normalize_0_36, plain,
    (! [A, B] : (disjoint(A, B) | ? [C] : in(C, set_intersection2(A, B)))),
    inference(conjunct, [], [normalize_0_35])).

fof(normalize_0_37, plain,
    (! [A, B] : (disjoint(A, B) | ? [C] : in(C, set_intersection2(A, B)))),
    inference(specialize, [], [normalize_0_36])).

fof(normalize_0_38, plain,
    (! [A, B] :
       (disjoint(A, B) |
        in(skolemFOFtoCNF_C_3(A, B), set_intersection2(A, B)))),
    inference(clausify, [], [normalize_0_37])).

fof(normalize_0_39, plain,
    (! [A, B, C] :
       (C != set_intersection2(A, B) <=>
        ? [D] : (~ in(D, C) <=> (in(D, A) & in(D, B))))),
    inference(canonicalize, [], [d3_xboole_0])).

fof(normalize_0_40, plain,
    (! [A, B, C] :
       (C != set_intersection2(A, B) <=>
        ? [D] : (~ in(D, C) <=> (in(D, A) & in(D, B))))),
    inference(specialize, [], [normalize_0_39])).

fof(normalize_0_41, plain,
    (! [A, B, C, D] :
       ((C != set_intersection2(A, B) | ~ in(D, C) | in(D, A)) &
        (C != set_intersection2(A, B) | ~ in(D, C) | in(D, B)) &
        (C = set_intersection2(A, B) | in(skolemFOFtoCNF_D_1(A, B, C), A) |
         in(skolemFOFtoCNF_D_1(A, B, C), C)) &
        (C = set_intersection2(A, B) | in(skolemFOFtoCNF_D_1(A, B, C), B) |
         in(skolemFOFtoCNF_D_1(A, B, C), C)) &
        (C != set_intersection2(A, B) | ~ in(D, A) | ~ in(D, B) |
         in(D, C)) &
        (~ in(skolemFOFtoCNF_D_1(A, B, C), A) |
         ~ in(skolemFOFtoCNF_D_1(A, B, C), B) |
         ~ in(skolemFOFtoCNF_D_1(A, B, C), C) |
         C = set_intersection2(A, B)))),
    inference(clausify, [], [normalize_0_40])).

fof(normalize_0_42, plain,
    (! [A, B, C, D] :
       (C != set_intersection2(A, B) | ~ in(D, C) | in(D, A))),
    inference(conjunct, [], [normalize_0_41])).

fof(normalize_0_43, plain,
    (! [A, B, C, D] :
       (C != set_intersection2(A, B) | ~ in(D, C) | in(D, B))),
    inference(conjunct, [], [normalize_0_41])).

fof(normalize_0_44, plain,
    (subset(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(conjunct, [], [normalize_0_13])).

fof(normalize_0_45, plain,
    (~ disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4)),
    inference(conjunct, [], [normalize_0_13])).

cnf(refute_0_0, plain, (~ disjoint(A, B) | disjoint(B, A)),
    inference(canonicalize, [], [normalize_0_1])).

cnf(refute_0_1, plain,
    (~ disjoint(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2) |
     disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_0 :
               [bind(A, $fot(skolemFOFtoCNF_C_4)),
                bind(B, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_2, plain,
    (C != set_difference(A, B) | ~ in(D, B) | ~ in(D, C)),
    inference(canonicalize, [], [normalize_0_5])).

cnf(refute_0_3, plain,
    (set_difference(A, B) != set_difference(A, B) | ~ in(D, B) |
     ~ in(D, set_difference(A, B))),
    inference(subst, [],
              [refute_0_2 : [bind(C, $fot(set_difference(A, B)))]])).

cnf(refute_0_4, plain, (set_difference(A, B) = set_difference(A, B)),
    introduced(tautology, [refl, [$fot(set_difference(A, B))]])).

cnf(refute_0_5, plain, (~ in(D, B) | ~ in(D, set_difference(A, B))),
    inference(resolve,
              [$cnf($equal(set_difference(A, B), set_difference(A, B)))],
              [refute_0_4, refute_0_3])).

cnf(refute_0_6, plain,
    (~ in(X_284, set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     ~ in(X_284, skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_5 :
               [bind(A, $fot(skolemFOFtoCNF_B_1)),
                bind(B, $fot(skolemFOFtoCNF_C_4)),
                bind(D, $fot(X_284))]])).

cnf(refute_0_7, plain,
    (set_union2(A, set_difference(B, A)) = set_union2(A, B)),
    inference(canonicalize, [], [normalize_0_7])).

cnf(refute_0_8, plain,
    (set_union2(set_difference(B, X_63),
       set_difference(B, set_difference(B, X_63))) =
     set_union2(set_difference(B, X_63), B)),
    inference(subst, [],
              [refute_0_7 : [bind(A, $fot(set_difference(B, X_63)))]])).

cnf(refute_0_9, plain,
    (set_difference(A, set_difference(A, B)) = set_intersection2(A, B)),
    inference(canonicalize, [], [normalize_0_9])).

cnf(refute_0_10, plain,
    (set_difference(B, set_difference(B, X_63)) =
     set_intersection2(B, X_63)),
    inference(subst, [],
              [refute_0_9 : [bind(A, $fot(B)), bind(B, $fot(X_63))]])).

cnf(refute_0_11, plain,
    (set_difference(B, set_difference(B, X_63)) !=
     set_intersection2(B, X_63) |
     set_union2(set_difference(B, X_63),
       set_difference(B, set_difference(B, X_63))) !=
     set_union2(set_difference(B, X_63), B) |
     set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(set_difference(B, X_63), B)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(B, X_63),
                          set_difference(B, set_difference(B, X_63))),
                        set_union2(set_difference(B, X_63), B))), [0, 1],
                 $fot(set_intersection2(B, X_63))]])).

cnf(refute_0_12, plain,
    (set_union2(set_difference(B, X_63),
       set_difference(B, set_difference(B, X_63))) !=
     set_union2(set_difference(B, X_63), B) |
     set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(set_difference(B, X_63), B)),
    inference(resolve,
              [$cnf($equal(set_difference(B, set_difference(B, X_63)),
                      set_intersection2(B, X_63)))],
              [refute_0_10, refute_0_11])).

cnf(refute_0_13, plain,
    (set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(set_difference(B, X_63), B)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(B, X_63),
                        set_difference(B, set_difference(B, X_63))),
                      set_union2(set_difference(B, X_63), B)))],
              [refute_0_8, refute_0_12])).

cnf(refute_0_14, plain, (set_union2(A, B) = set_union2(B, A)),
    inference(canonicalize, [], [normalize_0_11])).

cnf(refute_0_15, plain, (X = X), introduced(tautology, [refl, [$fot(X)]])).

cnf(refute_0_16, plain, (X != X | X != Y | Y = X),
    introduced(tautology, [equality, [$cnf($equal(X, X)), [0], $fot(Y)]])).

cnf(refute_0_17, plain, (X != Y | Y = X),
    inference(resolve, [$cnf($equal(X, X))], [refute_0_15, refute_0_16])).

cnf(refute_0_18, plain,
    (set_union2(A, B) != set_union2(B, A) |
     set_union2(B, A) = set_union2(A, B)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(A, B))),
                bind(Y, $fot(set_union2(B, A)))]])).

cnf(refute_0_19, plain, (set_union2(B, A) = set_union2(A, B)),
    inference(resolve, [$cnf($equal(set_union2(A, B), set_union2(B, A)))],
              [refute_0_14, refute_0_18])).

cnf(refute_0_20, plain,
    (set_union2(set_difference(B, X_63), B) =
     set_union2(B, set_difference(B, X_63))),
    inference(subst, [],
              [refute_0_19 :
               [bind(A, $fot(B)),
                bind(B, $fot(set_difference(B, X_63)))]])).

cnf(refute_0_21, plain,
    (set_union2(set_difference(B, X_63), B) !=
     set_union2(B, set_difference(B, X_63)) |
     set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) !=
     set_union2(set_difference(B, X_63), B) |
     set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(B, set_difference(B, X_63))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(B, X_63),
                          set_intersection2(B, X_63)),
                        set_union2(set_difference(B, X_63), B))), [1],
                 $fot(set_union2(B, set_difference(B, X_63)))]])).

cnf(refute_0_22, plain,
    (set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) !=
     set_union2(set_difference(B, X_63), B) |
     set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(B, set_difference(B, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(B, X_63), B),
                      set_union2(B, set_difference(B, X_63))))],
              [refute_0_20, refute_0_21])).

cnf(refute_0_23, plain,
    (set_union2(set_difference(B, X_63), set_intersection2(B, X_63)) =
     set_union2(B, set_difference(B, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(B, X_63),
                        set_intersection2(B, X_63)),
                      set_union2(set_difference(B, X_63), B)))],
              [refute_0_13, refute_0_22])).

cnf(refute_0_24, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(subst, [],
              [refute_0_23 :
               [bind(B, $fot(skolemFOFtoCNF_B_1)),
                bind(X_63, $fot(skolemFOFtoCNF_C_4))]])).

cnf(refute_0_25, plain, (disjoint(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(canonicalize, [], [normalize_0_14])).

cnf(refute_0_26, plain,
    (~ disjoint(A, B) | set_intersection2(A, B) = empty_set),
    inference(canonicalize, [], [normalize_0_18])).

cnf(refute_0_27, plain,
    (~ disjoint(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) |
     set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     empty_set),
    inference(subst, [],
              [refute_0_26 :
               [bind(A, $fot(skolemFOFtoCNF_B_1)),
                bind(B, $fot(skolemFOFtoCNF_C_4))]])).

cnf(refute_0_28, plain,
    (set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     empty_set),
    inference(resolve,
              [$cnf(disjoint(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))],
              [refute_0_25, refute_0_27])).

cnf(refute_0_29, plain,
    (set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) !=
     empty_set |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4),
                          set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4)),
                        set_union2(skolemFOFtoCNF_B_1,
                          set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4)))), [0, 1],
                 $fot(empty_set)]])).

cnf(refute_0_30, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(resolve,
              [$cnf($equal(set_intersection2(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4), empty_set))],
              [refute_0_28, refute_0_29])).

cnf(refute_0_31, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4),
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4)),
                      set_union2(skolemFOFtoCNF_B_1,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4))))],
              [refute_0_24, refute_0_30])).

cnf(refute_0_32, plain, (set_union2(A, empty_set) = A),
    inference(canonicalize, [], [normalize_0_20])).

cnf(refute_0_33, plain,
    (set_union2(A, empty_set) = set_union2(empty_set, A)),
    inference(subst, [], [refute_0_14 : [bind(B, $fot(empty_set))]])).

cnf(refute_0_34, plain,
    (set_union2(A, empty_set) != A |
     set_union2(A, empty_set) != set_union2(empty_set, A) |
     set_union2(empty_set, A) = A),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(A, empty_set), A)), [0],
                 $fot(set_union2(empty_set, A))]])).

cnf(refute_0_35, plain,
    (set_union2(A, empty_set) != A | set_union2(empty_set, A) = A),
    inference(resolve,
              [$cnf($equal(set_union2(A, empty_set),
                      set_union2(empty_set, A)))],
              [refute_0_33, refute_0_34])).

cnf(refute_0_36, plain, (set_union2(empty_set, A) = A),
    inference(resolve, [$cnf($equal(set_union2(A, empty_set), A))],
              [refute_0_32, refute_0_35])).

cnf(refute_0_37, plain,
    (set_union2(empty_set,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) =
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_36 :
               [bind(A,
                     $fot(set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4)))]])).

cnf(refute_0_38, plain,
    (set_difference(set_union2(A, X_63),
       set_difference(set_union2(A, X_63), X_63)) =
     set_intersection2(set_union2(A, X_63), X_63)),
    inference(subst, [],
              [refute_0_9 :
               [bind(A, $fot(set_union2(A, X_63))),
                bind(B, $fot(X_63))]])).

cnf(refute_0_39, plain,
    (set_difference(set_union2(A, B), B) = set_difference(A, B)),
    inference(canonicalize, [], [normalize_0_22])).

cnf(refute_0_40, plain,
    (set_difference(set_union2(A, X_63), X_63) = set_difference(A, X_63)),
    inference(subst, [], [refute_0_39 : [bind(B, $fot(X_63))]])).

cnf(refute_0_41, plain,
    (set_difference(set_union2(A, X_63), X_63) != set_difference(A, X_63) |
     set_difference(set_union2(A, X_63),
       set_difference(set_union2(A, X_63), X_63)) !=
     set_intersection2(set_union2(A, X_63), X_63) |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(set_union2(A, X_63), X_63)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(A, X_63),
                          set_difference(set_union2(A, X_63), X_63)),
                        set_intersection2(set_union2(A, X_63), X_63))),
                 [0, 1], $fot(set_difference(A, X_63))]])).

cnf(refute_0_42, plain,
    (set_difference(set_union2(A, X_63),
       set_difference(set_union2(A, X_63), X_63)) !=
     set_intersection2(set_union2(A, X_63), X_63) |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(set_union2(A, X_63), X_63)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_63), X_63),
                      set_difference(A, X_63)))],
              [refute_0_40, refute_0_41])).

cnf(refute_0_43, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(set_union2(A, X_63), X_63)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_63),
                        set_difference(set_union2(A, X_63), X_63)),
                      set_intersection2(set_union2(A, X_63), X_63)))],
              [refute_0_38, refute_0_42])).

cnf(refute_0_44, plain,
    (set_intersection2(A, B) = set_intersection2(B, A)),
    inference(canonicalize, [], [normalize_0_24])).

cnf(refute_0_45, plain,
    (set_intersection2(A, B) != set_intersection2(B, A) |
     set_intersection2(B, A) = set_intersection2(A, B)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_intersection2(A, B))),
                bind(Y, $fot(set_intersection2(B, A)))]])).

cnf(refute_0_46, plain,
    (set_intersection2(B, A) = set_intersection2(A, B)),
    inference(resolve,
              [$cnf($equal(set_intersection2(A, B),
                      set_intersection2(B, A)))],
              [refute_0_44, refute_0_45])).

cnf(refute_0_47, plain,
    (set_intersection2(set_union2(A, X_63), X_63) =
     set_intersection2(X_63, set_union2(A, X_63))),
    inference(subst, [],
              [refute_0_46 :
               [bind(A, $fot(X_63)),
                bind(B, $fot(set_union2(A, X_63)))]])).

cnf(refute_0_48, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) !=
     set_intersection2(set_union2(A, X_63), X_63) |
     set_intersection2(set_union2(A, X_63), X_63) !=
     set_intersection2(X_63, set_union2(A, X_63)) |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(X_63, set_union2(A, X_63))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(set_union2(A, X_63),
                          set_difference(A, X_63)),
                        set_intersection2(X_63, set_union2(A, X_63)))),
                 [0],
                 $fot(set_intersection2(set_union2(A, X_63), X_63))]])).

cnf(refute_0_49, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) !=
     set_intersection2(set_union2(A, X_63), X_63) |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(X_63, set_union2(A, X_63))),
    inference(resolve,
              [$cnf($equal(set_intersection2(set_union2(A, X_63), X_63),
                      set_intersection2(X_63, set_union2(A, X_63))))],
              [refute_0_47, refute_0_48])).

cnf(refute_0_50, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) =
     set_intersection2(X_63, set_union2(A, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_63),
                        set_difference(A, X_63)),
                      set_intersection2(set_union2(A, X_63), X_63)))],
              [refute_0_43, refute_0_49])).

cnf(refute_0_51, plain,
    (set_difference(A, set_difference(A, set_union2(X_71, A))) =
     set_intersection2(A, set_union2(X_71, A))),
    inference(subst, [],
              [refute_0_9 : [bind(B, $fot(set_union2(X_71, A)))]])).

cnf(refute_0_52, plain,
    (set_difference(set_union2(set_union2(X_70, X_69),
         set_union2(X_69, X_70)), set_union2(X_69, X_70)) =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    inference(subst, [],
              [refute_0_39 :
               [bind(A, $fot(set_union2(X_70, X_69))),
                bind(B, $fot(set_union2(X_69, X_70)))]])).

cnf(refute_0_53, plain,
    (set_union2(X_58, set_difference(set_union2(A, X_58), X_58)) =
     set_union2(X_58, set_union2(A, X_58))),
    inference(subst, [],
              [refute_0_7 :
               [bind(A, $fot(X_58)),
                bind(B, $fot(set_union2(A, X_58)))]])).

cnf(refute_0_54, plain,
    (set_difference(set_union2(A, X_58), X_58) = set_difference(A, X_58)),
    inference(subst, [], [refute_0_39 : [bind(B, $fot(X_58))]])).

cnf(refute_0_55, plain,
    (set_difference(set_union2(A, X_58), X_58) != set_difference(A, X_58) |
     set_union2(X_58, set_difference(set_union2(A, X_58), X_58)) !=
     set_union2(X_58, set_union2(A, X_58)) |
     set_union2(X_58, set_difference(A, X_58)) =
     set_union2(X_58, set_union2(A, X_58))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(X_58,
                          set_difference(set_union2(A, X_58), X_58)),
                        set_union2(X_58, set_union2(A, X_58)))), [0, 1],
                 $fot(set_difference(A, X_58))]])).

cnf(refute_0_56, plain,
    (set_union2(X_58, set_difference(set_union2(A, X_58), X_58)) !=
     set_union2(X_58, set_union2(A, X_58)) |
     set_union2(X_58, set_difference(A, X_58)) =
     set_union2(X_58, set_union2(A, X_58))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_58), X_58),
                      set_difference(A, X_58)))],
              [refute_0_54, refute_0_55])).

cnf(refute_0_57, plain,
    (set_union2(X_58, set_difference(A, X_58)) =
     set_union2(X_58, set_union2(A, X_58))),
    inference(resolve,
              [$cnf($equal(set_union2(X_58,
                        set_difference(set_union2(A, X_58), X_58)),
                      set_union2(X_58, set_union2(A, X_58))))],
              [refute_0_53, refute_0_56])).

cnf(refute_0_58, plain,
    (set_union2(X_58, set_difference(A, X_58)) = set_union2(X_58, A)),
    inference(subst, [],
              [refute_0_7 : [bind(A, $fot(X_58)), bind(B, $fot(A))]])).

cnf(refute_0_59, plain,
    (set_union2(X_58, set_difference(A, X_58)) != set_union2(X_58, A) |
     set_union2(X_58, set_difference(A, X_58)) !=
     set_union2(X_58, set_union2(A, X_58)) |
     set_union2(X_58, A) = set_union2(X_58, set_union2(A, X_58))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(X_58, set_difference(A, X_58)),
                        set_union2(X_58, set_union2(A, X_58)))), [0],
                 $fot(set_union2(X_58, A))]])).

cnf(refute_0_60, plain,
    (set_union2(X_58, set_difference(A, X_58)) !=
     set_union2(X_58, set_union2(A, X_58)) |
     set_union2(X_58, A) = set_union2(X_58, set_union2(A, X_58))),
    inference(resolve,
              [$cnf($equal(set_union2(X_58, set_difference(A, X_58)),
                      set_union2(X_58, A)))], [refute_0_58, refute_0_59])).

cnf(refute_0_61, plain,
    (set_union2(X_58, A) = set_union2(X_58, set_union2(A, X_58))),
    inference(resolve,
              [$cnf($equal(set_union2(X_58, set_difference(A, X_58)),
                      set_union2(X_58, set_union2(A, X_58))))],
              [refute_0_57, refute_0_60])).

cnf(refute_0_62, plain,
    (set_union2(set_union2(X_60, A), A) =
     set_union2(set_union2(X_60, A), set_union2(A, set_union2(X_60, A)))),
    inference(subst, [],
              [refute_0_61 : [bind(X_58, $fot(set_union2(X_60, A)))]])).

cnf(refute_0_63, plain,
    (set_union2(A, X_60) = set_union2(A, set_union2(X_60, A))),
    inference(subst, [],
              [refute_0_61 : [bind(A, $fot(X_60)), bind(X_58, $fot(A))]])).

cnf(refute_0_64, plain,
    (set_union2(A, X_60) != set_union2(A, set_union2(X_60, A)) |
     set_union2(A, set_union2(X_60, A)) = set_union2(A, X_60)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(A, X_60))),
                bind(Y, $fot(set_union2(A, set_union2(X_60, A))))]])).

cnf(refute_0_65, plain,
    (set_union2(A, set_union2(X_60, A)) = set_union2(A, X_60)),
    inference(resolve,
              [$cnf($equal(set_union2(A, X_60),
                      set_union2(A, set_union2(X_60, A))))],
              [refute_0_63, refute_0_64])).

cnf(refute_0_66, plain,
    (set_union2(A, set_union2(X_60, A)) != set_union2(A, X_60) |
     set_union2(set_union2(X_60, A), A) !=
     set_union2(set_union2(X_60, A), set_union2(A, set_union2(X_60, A))) |
     set_union2(set_union2(X_60, A), A) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_union2(X_60, A), A),
                        set_union2(set_union2(X_60, A),
                          set_union2(A, set_union2(X_60, A))))), [1, 1],
                 $fot(set_union2(A, X_60))]])).

cnf(refute_0_67, plain,
    (set_union2(set_union2(X_60, A), A) !=
     set_union2(set_union2(X_60, A), set_union2(A, set_union2(X_60, A))) |
     set_union2(set_union2(X_60, A), A) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    inference(resolve,
              [$cnf($equal(set_union2(A, set_union2(X_60, A)),
                      set_union2(A, X_60)))], [refute_0_65, refute_0_66])).

cnf(refute_0_68, plain,
    (set_union2(set_union2(X_60, A), A) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(X_60, A), A),
                      set_union2(set_union2(X_60, A),
                        set_union2(A, set_union2(X_60, A)))))],
              [refute_0_62, refute_0_67])).

cnf(refute_0_69, plain,
    (set_union2(set_union2(X_60, A), A) =
     set_union2(A, set_union2(X_60, A))),
    inference(subst, [],
              [refute_0_19 : [bind(B, $fot(set_union2(X_60, A)))]])).

cnf(refute_0_70, plain, (Y != X | Y != Z | X = Z),
    introduced(tautology, [equality, [$cnf($equal(Y, Z)), [0], $fot(X)]])).

cnf(refute_0_71, plain, (X != Y | Y != Z | X = Z),
    inference(resolve, [$cnf($equal(Y, X))], [refute_0_17, refute_0_70])).

cnf(refute_0_72, plain,
    (set_union2(A, set_union2(X_60, A)) != set_union2(A, X_60) |
     set_union2(set_union2(X_60, A), A) !=
     set_union2(A, set_union2(X_60, A)) |
     set_union2(set_union2(X_60, A), A) = set_union2(A, X_60)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X, $fot(set_union2(set_union2(X_60, A), A))),
                bind(Y, $fot(set_union2(A, set_union2(X_60, A)))),
                bind(Z, $fot(set_union2(A, X_60)))]])).

cnf(refute_0_73, plain,
    (set_union2(A, set_union2(X_60, A)) != set_union2(A, X_60) |
     set_union2(set_union2(X_60, A), A) = set_union2(A, X_60)),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(X_60, A), A),
                      set_union2(A, set_union2(X_60, A))))],
              [refute_0_69, refute_0_72])).

cnf(refute_0_74, plain,
    (set_union2(set_union2(X_60, A), A) = set_union2(A, X_60)),
    inference(resolve,
              [$cnf($equal(set_union2(A, set_union2(X_60, A)),
                      set_union2(A, X_60)))], [refute_0_65, refute_0_73])).

cnf(refute_0_75, plain,
    (set_union2(set_union2(X_60, A), A) != set_union2(A, X_60) |
     set_union2(set_union2(X_60, A), A) !=
     set_union2(set_union2(X_60, A), set_union2(A, X_60)) |
     set_union2(A, X_60) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_union2(X_60, A), A),
                        set_union2(set_union2(X_60, A),
                          set_union2(A, X_60)))), [0],
                 $fot(set_union2(A, X_60))]])).

cnf(refute_0_76, plain,
    (set_union2(set_union2(X_60, A), A) !=
     set_union2(set_union2(X_60, A), set_union2(A, X_60)) |
     set_union2(A, X_60) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(X_60, A), A),
                      set_union2(A, X_60)))], [refute_0_74, refute_0_75])).

cnf(refute_0_77, plain,
    (set_union2(A, X_60) =
     set_union2(set_union2(X_60, A), set_union2(A, X_60))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(X_60, A), A),
                      set_union2(set_union2(X_60, A),
                        set_union2(A, X_60))))],
              [refute_0_68, refute_0_76])).

cnf(refute_0_78, plain,
    (set_union2(X_69, X_70) =
     set_union2(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    inference(subst, [],
              [refute_0_77 :
               [bind(A, $fot(X_69)), bind(X_60, $fot(X_70))]])).

cnf(refute_0_79, plain,
    (set_union2(X_69, X_70) !=
     set_union2(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     set_union2(set_union2(X_70, X_69), set_union2(X_69, X_70)) =
     set_union2(X_69, X_70)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(X_69, X_70))),
                bind(Y,
                     $fot(set_union2(set_union2(X_70, X_69),
                            set_union2(X_69, X_70))))]])).

cnf(refute_0_80, plain,
    (set_union2(set_union2(X_70, X_69), set_union2(X_69, X_70)) =
     set_union2(X_69, X_70)),
    inference(resolve,
              [$cnf($equal(set_union2(X_69, X_70),
                      set_union2(set_union2(X_70, X_69),
                        set_union2(X_69, X_70))))],
              [refute_0_78, refute_0_79])).

cnf(refute_0_81, plain,
    (set_difference(set_union2(set_union2(X_70, X_69),
         set_union2(X_69, X_70)), set_union2(X_69, X_70)) !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     set_union2(set_union2(X_70, X_69), set_union2(X_69, X_70)) !=
     set_union2(X_69, X_70) |
     set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(set_union2(X_70,
                              X_69), set_union2(X_69, X_70)),
                          set_union2(X_69, X_70)),
                        set_difference(set_union2(X_70, X_69),
                          set_union2(X_69, X_70)))), [0, 0],
                 $fot(set_union2(X_69, X_70))]])).

cnf(refute_0_82, plain,
    (set_difference(set_union2(set_union2(X_70, X_69),
         set_union2(X_69, X_70)), set_union2(X_69, X_70)) !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(X_70, X_69),
                        set_union2(X_69, X_70)), set_union2(X_69, X_70)))],
              [refute_0_80, refute_0_81])).

cnf(refute_0_83, plain,
    (set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_union2(X_70,
                            X_69), set_union2(X_69, X_70)),
                        set_union2(X_69, X_70)),
                      set_difference(set_union2(X_70, X_69),
                        set_union2(X_69, X_70))))],
              [refute_0_52, refute_0_82])).

cnf(refute_0_84, plain,
    (set_difference(set_union2(empty_set, X_51), X_51) =
     set_difference(empty_set, X_51)),
    inference(subst, [],
              [refute_0_39 :
               [bind(A, $fot(empty_set)), bind(B, $fot(X_51))]])).

cnf(refute_0_85, plain, (set_union2(empty_set, X_51) = X_51),
    inference(subst, [], [refute_0_36 : [bind(A, $fot(X_51))]])).

cnf(refute_0_86, plain,
    (set_difference(set_union2(empty_set, X_51), X_51) !=
     set_difference(empty_set, X_51) |
     set_union2(empty_set, X_51) != X_51 |
     set_difference(X_51, X_51) = set_difference(empty_set, X_51)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(empty_set, X_51),
                          X_51), set_difference(empty_set, X_51))), [0, 0],
                 $fot(X_51)]])).

cnf(refute_0_87, plain,
    (set_difference(set_union2(empty_set, X_51), X_51) !=
     set_difference(empty_set, X_51) |
     set_difference(X_51, X_51) = set_difference(empty_set, X_51)),
    inference(resolve, [$cnf($equal(set_union2(empty_set, X_51), X_51))],
              [refute_0_85, refute_0_86])).

cnf(refute_0_88, plain,
    (set_difference(X_51, X_51) = set_difference(empty_set, X_51)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(empty_set, X_51),
                        X_51), set_difference(empty_set, X_51)))],
              [refute_0_84, refute_0_87])).

cnf(refute_0_89, plain, (set_difference(empty_set, A) = empty_set),
    inference(canonicalize, [], [normalize_0_26])).

cnf(refute_0_90, plain, (set_difference(empty_set, X_51) = empty_set),
    inference(subst, [], [refute_0_89 : [bind(A, $fot(X_51))]])).

cnf(refute_0_91, plain,
    (set_difference(X_51, X_51) != set_difference(empty_set, X_51) |
     set_difference(empty_set, X_51) != empty_set |
     set_difference(X_51, X_51) = empty_set),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(X_51, X_51), empty_set)),
                 [0], $fot(set_difference(empty_set, X_51))]])).

cnf(refute_0_92, plain,
    (set_difference(X_51, X_51) != set_difference(empty_set, X_51) |
     set_difference(X_51, X_51) = empty_set),
    inference(resolve,
              [$cnf($equal(set_difference(empty_set, X_51), empty_set))],
              [refute_0_90, refute_0_91])).

cnf(refute_0_93, plain, (set_difference(X_51, X_51) = empty_set),
    inference(resolve,
              [$cnf($equal(set_difference(X_51, X_51),
                      set_difference(empty_set, X_51)))],
              [refute_0_88, refute_0_92])).

cnf(refute_0_94, plain,
    (set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) =
     empty_set),
    inference(subst, [],
              [refute_0_93 : [bind(X_51, $fot(set_union2(X_69, X_70)))]])).

cnf(refute_0_95, plain,
    (set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) !=
     empty_set |
     set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     empty_set =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(X_69, X_70),
                          set_union2(X_69, X_70)),
                        set_difference(set_union2(X_70, X_69),
                          set_union2(X_69, X_70)))), [0],
                 $fot(empty_set)]])).

cnf(refute_0_96, plain,
    (set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     empty_set =
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_69, X_70),
                        set_union2(X_69, X_70)), empty_set))],
              [refute_0_94, refute_0_95])).

cnf(refute_0_97, plain,
    (set_difference(set_union2(A, set_union2(X_60, A)),
       set_union2(X_60, A)) = set_difference(A, set_union2(X_60, A))),
    inference(subst, [],
              [refute_0_39 : [bind(B, $fot(set_union2(X_60, A)))]])).

cnf(refute_0_98, plain,
    (set_difference(set_union2(A, set_union2(X_60, A)),
       set_union2(X_60, A)) != set_difference(A, set_union2(X_60, A)) |
     set_union2(A, set_union2(X_60, A)) != set_union2(A, X_60) |
     set_difference(set_union2(A, X_60), set_union2(X_60, A)) =
     set_difference(A, set_union2(X_60, A))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(A,
                            set_union2(X_60, A)), set_union2(X_60, A)),
                        set_difference(A, set_union2(X_60, A)))), [0, 0],
                 $fot(set_union2(A, X_60))]])).

cnf(refute_0_99, plain,
    (set_difference(set_union2(A, set_union2(X_60, A)),
       set_union2(X_60, A)) != set_difference(A, set_union2(X_60, A)) |
     set_difference(set_union2(A, X_60), set_union2(X_60, A)) =
     set_difference(A, set_union2(X_60, A))),
    inference(resolve,
              [$cnf($equal(set_union2(A, set_union2(X_60, A)),
                      set_union2(A, X_60)))], [refute_0_65, refute_0_98])).

cnf(refute_0_100, plain,
    (set_difference(set_union2(A, X_60), set_union2(X_60, A)) =
     set_difference(A, set_union2(X_60, A))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A,
                          set_union2(X_60, A)), set_union2(X_60, A)),
                      set_difference(A, set_union2(X_60, A))))],
              [refute_0_97, refute_0_99])).

cnf(refute_0_101, plain,
    (set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) =
     set_difference(X_70, set_union2(X_69, X_70))),
    inference(subst, [],
              [refute_0_100 :
               [bind(A, $fot(X_70)), bind(X_60, $fot(X_69))]])).

cnf(refute_0_102, plain,
    (empty_set !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) !=
     set_difference(X_70, set_union2(X_69, X_70)) |
     empty_set = set_difference(X_70, set_union2(X_69, X_70))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(empty_set,
                        set_difference(X_70, set_union2(X_69, X_70)))),
                 [0],
                 $fot(set_difference(set_union2(X_70, X_69),
                        set_union2(X_69, X_70)))]])).

cnf(refute_0_103, plain,
    (empty_set !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     empty_set = set_difference(X_70, set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_70, X_69),
                        set_union2(X_69, X_70)),
                      set_difference(X_70, set_union2(X_69, X_70))))],
              [refute_0_101, refute_0_102])).

cnf(refute_0_104, plain,
    (set_difference(set_union2(X_69, X_70), set_union2(X_69, X_70)) !=
     set_difference(set_union2(X_70, X_69), set_union2(X_69, X_70)) |
     empty_set = set_difference(X_70, set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(empty_set,
                      set_difference(set_union2(X_70, X_69),
                        set_union2(X_69, X_70))))],
              [refute_0_96, refute_0_103])).

cnf(refute_0_105, plain,
    (empty_set = set_difference(X_70, set_union2(X_69, X_70))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_69, X_70),
                        set_union2(X_69, X_70)),
                      set_difference(set_union2(X_70, X_69),
                        set_union2(X_69, X_70))))],
              [refute_0_83, refute_0_104])).

cnf(refute_0_106, plain,
    (empty_set = set_difference(A, set_union2(X_71, A))),
    inference(subst, [],
              [refute_0_105 :
               [bind(X_69, $fot(X_71)), bind(X_70, $fot(A))]])).

cnf(refute_0_107, plain,
    (empty_set != set_difference(A, set_union2(X_71, A)) |
     set_difference(A, set_union2(X_71, A)) = empty_set),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(empty_set)),
                bind(Y, $fot(set_difference(A, set_union2(X_71, A))))]])).

cnf(refute_0_108, plain,
    (set_difference(A, set_union2(X_71, A)) = empty_set),
    inference(resolve,
              [$cnf($equal(empty_set,
                      set_difference(A, set_union2(X_71, A))))],
              [refute_0_106, refute_0_107])).

cnf(refute_0_109, plain,
    (set_difference(A, set_difference(A, set_union2(X_71, A))) !=
     set_intersection2(A, set_union2(X_71, A)) |
     set_difference(A, set_union2(X_71, A)) != empty_set |
     set_difference(A, empty_set) =
     set_intersection2(A, set_union2(X_71, A))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(A,
                          set_difference(A, set_union2(X_71, A))),
                        set_intersection2(A, set_union2(X_71, A)))),
                 [0, 1], $fot(empty_set)]])).

cnf(refute_0_110, plain,
    (set_difference(A, set_difference(A, set_union2(X_71, A))) !=
     set_intersection2(A, set_union2(X_71, A)) |
     set_difference(A, empty_set) =
     set_intersection2(A, set_union2(X_71, A))),
    inference(resolve,
              [$cnf($equal(set_difference(A, set_union2(X_71, A)),
                      empty_set))], [refute_0_108, refute_0_109])).

cnf(refute_0_111, plain,
    (set_difference(A, empty_set) =
     set_intersection2(A, set_union2(X_71, A))),
    inference(resolve,
              [$cnf($equal(set_difference(A,
                        set_difference(A, set_union2(X_71, A))),
                      set_intersection2(A, set_union2(X_71, A))))],
              [refute_0_51, refute_0_110])).

cnf(refute_0_112, plain, (set_difference(A, empty_set) = A),
    inference(canonicalize, [], [normalize_0_28])).

cnf(refute_0_113, plain,
    (set_difference(A, empty_set) != A |
     set_difference(A, empty_set) !=
     set_intersection2(A, set_union2(X_71, A)) |
     A = set_intersection2(A, set_union2(X_71, A))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(A, empty_set),
                        set_intersection2(A, set_union2(X_71, A)))), [0],
                 $fot(A)]])).

cnf(refute_0_114, plain,
    (set_difference(A, empty_set) !=
     set_intersection2(A, set_union2(X_71, A)) |
     A = set_intersection2(A, set_union2(X_71, A))),
    inference(resolve, [$cnf($equal(set_difference(A, empty_set), A))],
              [refute_0_112, refute_0_113])).

cnf(refute_0_115, plain, (A = set_intersection2(A, set_union2(X_71, A))),
    inference(resolve,
              [$cnf($equal(set_difference(A, empty_set),
                      set_intersection2(A, set_union2(X_71, A))))],
              [refute_0_111, refute_0_114])).

cnf(refute_0_116, plain,
    (A != set_intersection2(A, set_union2(X_71, A)) |
     set_intersection2(A, set_union2(X_71, A)) = A),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(A)),
                bind(Y,
                     $fot(set_intersection2(A, set_union2(X_71, A))))]])).

cnf(refute_0_117, plain, (set_intersection2(A, set_union2(X_71, A)) = A),
    inference(resolve,
              [$cnf($equal(A, set_intersection2(A, set_union2(X_71, A))))],
              [refute_0_115, refute_0_116])).

cnf(refute_0_118, plain,
    (set_intersection2(X_63, set_union2(A, X_63)) = X_63),
    inference(subst, [],
              [refute_0_117 :
               [bind(A, $fot(X_63)), bind(X_71, $fot(A))]])).

cnf(refute_0_119, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) !=
     set_intersection2(X_63, set_union2(A, X_63)) |
     set_intersection2(X_63, set_union2(A, X_63)) != X_63 |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) = X_63),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(set_union2(A, X_63),
                          set_difference(A, X_63)), X_63)), [0],
                 $fot(set_intersection2(X_63, set_union2(A, X_63)))]])).

cnf(refute_0_120, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) !=
     set_intersection2(X_63, set_union2(A, X_63)) |
     set_difference(set_union2(A, X_63), set_difference(A, X_63)) = X_63),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_63, set_union2(A, X_63)),
                      X_63))], [refute_0_118, refute_0_119])).

cnf(refute_0_121, plain,
    (set_difference(set_union2(A, X_63), set_difference(A, X_63)) = X_63),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_63),
                        set_difference(A, X_63)),
                      set_intersection2(X_63, set_union2(A, X_63))))],
              [refute_0_50, refute_0_120])).

cnf(refute_0_122, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_187, X_188)),
    inference(subst, [],
              [refute_0_121 :
               [bind(A, $fot(set_difference(X_187, X_188))),
                bind(X_63, $fot(set_union2(X_187, X_188)))]])).

cnf(refute_0_123, plain,
    (set_union2(set_difference(set_union2(A, X_182), X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    inference(subst, [],
              [refute_0_23 :
               [bind(B, $fot(set_union2(A, X_182))),
                bind(X_63, $fot(X_182))]])).

cnf(refute_0_124, plain,
    (set_difference(set_union2(A, X_182), X_182) =
     set_difference(A, X_182)),
    inference(subst, [], [refute_0_39 : [bind(B, $fot(X_182))]])).

cnf(refute_0_125, plain,
    (set_difference(set_union2(A, X_182), X_182) !=
     set_difference(A, X_182) |
     set_union2(set_difference(set_union2(A, X_182), X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(set_union2(A,
                              X_182), X_182),
                          set_intersection2(set_union2(A, X_182), X_182)),
                        set_union2(set_union2(A, X_182),
                          set_difference(set_union2(A, X_182), X_182)))),
                 [0, 0], $fot(set_difference(A, X_182))]])).

cnf(refute_0_126, plain,
    (set_union2(set_difference(set_union2(A, X_182), X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_182), X_182),
                      set_difference(A, X_182)))],
              [refute_0_124, refute_0_125])).

cnf(refute_0_127, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(set_union2(A, X_182),
                          X_182),
                        set_intersection2(set_union2(A, X_182), X_182)),
                      set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182))))],
              [refute_0_123, refute_0_126])).

cnf(refute_0_128, plain,
    (set_union2(X_182, set_difference(A, X_182)) = set_union2(X_182, A)),
    inference(subst, [],
              [refute_0_7 : [bind(A, $fot(X_182)), bind(B, $fot(A))]])).

cnf(refute_0_129, plain,
    (set_union2(set_difference(A, X_182), X_182) =
     set_union2(X_182, set_difference(A, X_182))),
    inference(subst, [],
              [refute_0_19 :
               [bind(A, $fot(X_182)),
                bind(B, $fot(set_difference(A, X_182)))]])).

cnf(refute_0_130, plain,
    (set_union2(X_182, set_difference(A, X_182)) != set_union2(X_182, A) |
     set_union2(set_difference(A, X_182), X_182) !=
     set_union2(X_182, set_difference(A, X_182)) |
     set_union2(set_difference(A, X_182), X_182) = set_union2(X_182, A)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X, $fot(set_union2(set_difference(A, X_182), X_182))),
                bind(Y, $fot(set_union2(X_182, set_difference(A, X_182)))),
                bind(Z, $fot(set_union2(X_182, A)))]])).

cnf(refute_0_131, plain,
    (set_union2(X_182, set_difference(A, X_182)) != set_union2(X_182, A) |
     set_union2(set_difference(A, X_182), X_182) = set_union2(X_182, A)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182), X_182),
                      set_union2(X_182, set_difference(A, X_182))))],
              [refute_0_129, refute_0_130])).

cnf(refute_0_132, plain,
    (set_union2(set_difference(A, X_182), X_182) = set_union2(X_182, A)),
    inference(resolve,
              [$cnf($equal(set_union2(X_182, set_difference(A, X_182)),
                      set_union2(X_182, A)))],
              [refute_0_128, refute_0_131])).

cnf(refute_0_133, plain,
    (set_intersection2(X_182, set_union2(A, X_182)) = X_182),
    inference(subst, [],
              [refute_0_117 :
               [bind(A, $fot(X_182)), bind(X_71, $fot(A))]])).

cnf(refute_0_134, plain,
    (set_intersection2(set_union2(A, X_182), X_182) =
     set_intersection2(X_182, set_union2(A, X_182))),
    inference(subst, [],
              [refute_0_46 :
               [bind(A, $fot(X_182)),
                bind(B, $fot(set_union2(A, X_182)))]])).

cnf(refute_0_135, plain,
    (set_intersection2(X_182, set_union2(A, X_182)) != X_182 |
     set_intersection2(set_union2(A, X_182), X_182) !=
     set_intersection2(X_182, set_union2(A, X_182)) |
     set_intersection2(set_union2(A, X_182), X_182) = X_182),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_intersection2(set_union2(A, X_182), X_182))),
                bind(Y,
                     $fot(set_intersection2(X_182, set_union2(A, X_182)))),
                bind(Z, $fot(X_182))]])).

cnf(refute_0_136, plain,
    (set_intersection2(X_182, set_union2(A, X_182)) != X_182 |
     set_intersection2(set_union2(A, X_182), X_182) = X_182),
    inference(resolve,
              [$cnf($equal(set_intersection2(set_union2(A, X_182), X_182),
                      set_intersection2(X_182, set_union2(A, X_182))))],
              [refute_0_134, refute_0_135])).

cnf(refute_0_137, plain,
    (set_intersection2(set_union2(A, X_182), X_182) = X_182),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_182, set_union2(A, X_182)),
                      X_182))], [refute_0_133, refute_0_136])).

cnf(refute_0_138, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182))),
    introduced(tautology,
               [refl,
                [$fot(set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182),
                          X_182)))]])).

cnf(refute_0_139, plain,
    (set_intersection2(set_union2(A, X_182), X_182) != X_182 |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), X_182)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(A, X_182),
                          set_intersection2(set_union2(A, X_182), X_182)),
                        set_union2(set_difference(A, X_182),
                          set_intersection2(set_union2(A, X_182),
                            X_182)))), [1, 1], $fot(X_182)]])).

cnf(refute_0_140, plain,
    (set_intersection2(set_union2(A, X_182), X_182) != X_182 |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), X_182)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182), X_182)),
                      set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182), X_182))))],
              [refute_0_138, refute_0_139])).

cnf(refute_0_141, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), X_182)),
    inference(resolve,
              [$cnf($equal(set_intersection2(set_union2(A, X_182), X_182),
                      X_182))], [refute_0_137, refute_0_140])).

cnf(refute_0_142, plain,
    (set_union2(set_difference(A, X_182), X_182) != set_union2(X_182, A) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_difference(A, X_182), X_182) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(X_182, A)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(A, X_182),
                            set_intersection2(set_union2(A, X_182),
                              X_182)))),
                bind(Y, $fot(set_union2(set_difference(A, X_182), X_182))),
                bind(Z, $fot(set_union2(X_182, A)))]])).

cnf(refute_0_143, plain,
    (set_union2(set_difference(A, X_182), X_182) != set_union2(X_182, A) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(X_182, A)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182), X_182)),
                      set_union2(set_difference(A, X_182), X_182)))],
              [refute_0_141, refute_0_142])).

cnf(refute_0_144, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) =
     set_union2(X_182, A)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182), X_182),
                      set_union2(X_182, A)))],
              [refute_0_132, refute_0_143])).

cnf(refute_0_145, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(X_182, A) |
     set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(X_182, A) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(A, X_182),
                          set_intersection2(set_union2(A, X_182), X_182)),
                        set_union2(set_union2(A, X_182),
                          set_difference(set_union2(A, X_182), X_182)))),
                 [0], $fot(set_union2(X_182, A))]])).

cnf(refute_0_146, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(X_182, A) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182), X_182)),
                      set_union2(X_182, A)))],
              [refute_0_144, refute_0_145])).

cnf(refute_0_147, plain,
    (set_union2(set_union2(A, X_182), set_difference(A, X_182)) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(subst, [],
              [refute_0_19 :
               [bind(A, $fot(set_difference(A, X_182))),
                bind(B, $fot(set_union2(A, X_182)))]])).

cnf(refute_0_148, plain,
    (set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182))),
    introduced(tautology,
               [refl,
                [$fot(set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182)))]])).

cnf(refute_0_149, plain,
    (set_difference(set_union2(A, X_182), X_182) !=
     set_difference(A, X_182) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182), set_difference(A, X_182))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_union2(A, X_182),
                          set_difference(set_union2(A, X_182), X_182)),
                        set_union2(set_union2(A, X_182),
                          set_difference(set_union2(A, X_182), X_182)))),
                 [1, 1], $fot(set_difference(A, X_182))]])).

cnf(refute_0_150, plain,
    (set_difference(set_union2(A, X_182), X_182) !=
     set_difference(A, X_182) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182), set_difference(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182)),
                      set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182))))],
              [refute_0_148, refute_0_149])).

cnf(refute_0_151, plain,
    (set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_union2(A, X_182), set_difference(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_182), X_182),
                      set_difference(A, X_182)))],
              [refute_0_124, refute_0_150])).

cnf(refute_0_152, plain,
    (set_union2(set_union2(A, X_182), set_difference(A, X_182)) !=
     set_union2(set_difference(A, X_182), set_union2(A, X_182)) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182), set_difference(A, X_182)) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_union2(A, X_182),
                            set_difference(set_union2(A, X_182), X_182)))),
                bind(Y,
                     $fot(set_union2(set_union2(A, X_182),
                            set_difference(A, X_182)))),
                bind(Z,
                     $fot(set_union2(set_difference(A, X_182),
                            set_union2(A, X_182))))]])).

cnf(refute_0_153, plain,
    (set_union2(set_union2(A, X_182), set_difference(A, X_182)) !=
     set_union2(set_difference(A, X_182), set_union2(A, X_182)) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182)),
                      set_union2(set_union2(A, X_182),
                        set_difference(A, X_182))))],
              [refute_0_151, refute_0_152])).

cnf(refute_0_154, plain,
    (set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(A, X_182),
                        set_difference(A, X_182)),
                      set_union2(set_difference(A, X_182),
                        set_union2(A, X_182))))],
              [refute_0_147, refute_0_153])).

cnf(refute_0_155, plain,
    (set_union2(X_182, A) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) !=
     set_union2(set_difference(A, X_182), set_union2(A, X_182)) |
     set_union2(X_182, A) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_union2(X_182, A),
                        set_union2(set_difference(A, X_182),
                          set_union2(A, X_182)))), [0],
                 $fot(set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182)))]])).

cnf(refute_0_156, plain,
    (set_union2(X_182, A) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(X_182, A) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182)),
                      set_union2(set_difference(A, X_182),
                        set_union2(A, X_182))))],
              [refute_0_154, refute_0_155])).

cnf(refute_0_157, plain,
    (set_union2(set_difference(A, X_182),
       set_intersection2(set_union2(A, X_182), X_182)) !=
     set_union2(set_union2(A, X_182),
       set_difference(set_union2(A, X_182), X_182)) |
     set_union2(X_182, A) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(X_182, A),
                      set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182))))],
              [refute_0_146, refute_0_156])).

cnf(refute_0_158, plain,
    (set_union2(X_182, A) =
     set_union2(set_difference(A, X_182), set_union2(A, X_182))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(A, X_182),
                        set_intersection2(set_union2(A, X_182), X_182)),
                      set_union2(set_union2(A, X_182),
                        set_difference(set_union2(A, X_182), X_182))))],
              [refute_0_127, refute_0_157])).

cnf(refute_0_159, plain,
    (set_union2(X_188, X_187) =
     set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188))),
    inference(subst, [],
              [refute_0_158 :
               [bind(A, $fot(X_187)), bind(X_182, $fot(X_188))]])).

cnf(refute_0_160, plain,
    (set_union2(X_188, X_187) !=
     set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188)) |
     set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188)) =
     set_union2(X_188, X_187)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(X_188, X_187))),
                bind(Y,
                     $fot(set_union2(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))))]])).

cnf(refute_0_161, plain,
    (set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188)) =
     set_union2(X_188, X_187)),
    inference(resolve,
              [$cnf($equal(set_union2(X_188, X_187),
                      set_union2(set_difference(X_187, X_188),
                        set_union2(X_187, X_188))))],
              [refute_0_159, refute_0_160])).

cnf(refute_0_162, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) != set_union2(X_187, X_188) |
     set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188)) !=
     set_union2(X_188, X_187) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_187, X_188)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(set_difference(X_187,
                              X_188), set_union2(X_187, X_188)),
                          set_difference(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))),
                        set_union2(X_187, X_188))), [0, 0],
                 $fot(set_union2(X_188, X_187))]])).

cnf(refute_0_163, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) != set_union2(X_187, X_188) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_187, X_188)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_187, X_188),
                        set_union2(X_187, X_188)),
                      set_union2(X_188, X_187)))],
              [refute_0_161, refute_0_162])).

cnf(refute_0_164, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_187, X_188)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_187,
                            X_188), set_union2(X_187, X_188)),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))),
                      set_union2(X_187, X_188)))],
              [refute_0_122, refute_0_163])).

cnf(refute_0_165, plain,
    (set_difference(set_union2(X_188, X_187), empty_set) =
     set_union2(X_188, X_187)),
    inference(subst, [],
              [refute_0_112 : [bind(A, $fot(set_union2(X_188, X_187)))]])).

cnf(refute_0_166, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)), set_union2(X_187, X_188)) =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    inference(subst, [],
              [refute_0_39 :
               [bind(A, $fot(set_difference(X_187, X_188))),
                bind(B, $fot(set_union2(X_187, X_188)))]])).

cnf(refute_0_167, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)), set_union2(X_187, X_188)) !=
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) |
     set_union2(set_difference(X_187, X_188), set_union2(X_187, X_188)) !=
     set_union2(X_188, X_187) |
     set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(set_difference(X_187,
                              X_188), set_union2(X_187, X_188)),
                          set_union2(X_187, X_188)),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188)))), [0, 0],
                 $fot(set_union2(X_188, X_187))]])).

cnf(refute_0_168, plain,
    (set_difference(set_union2(set_difference(X_187, X_188),
         set_union2(X_187, X_188)), set_union2(X_187, X_188)) !=
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) |
     set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_187, X_188),
                        set_union2(X_187, X_188)),
                      set_union2(X_188, X_187)))],
              [refute_0_161, refute_0_167])).

cnf(refute_0_169, plain,
    (set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_187,
                            X_188), set_union2(X_187, X_188)),
                        set_union2(X_187, X_188)),
                      set_difference(set_difference(X_187, X_188),
                        set_union2(X_187, X_188))))],
              [refute_0_166, refute_0_168])).

cnf(refute_0_170, plain,
    (empty_set != set_difference(X_70, set_union2(X_69, X_70)) |
     set_difference(X_70, set_union2(X_69, X_70)) = empty_set),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(empty_set)),
                bind(Y,
                     $fot(set_difference(X_70,
                            set_union2(X_69, X_70))))]])).

cnf(refute_0_171, plain,
    (set_difference(X_70, set_union2(X_69, X_70)) = empty_set),
    inference(resolve,
              [$cnf($equal(empty_set,
                      set_difference(X_70, set_union2(X_69, X_70))))],
              [refute_0_105, refute_0_170])).

cnf(refute_0_172, plain,
    (set_difference(A, set_union2(X_60, A)) = empty_set),
    inference(subst, [],
              [refute_0_171 :
               [bind(X_69, $fot(X_60)), bind(X_70, $fot(A))]])).

cnf(refute_0_173, plain,
    (set_difference(A, set_union2(X_60, A)) != empty_set |
     set_difference(set_union2(A, X_60), set_union2(X_60, A)) !=
     set_difference(A, set_union2(X_60, A)) |
     set_difference(set_union2(A, X_60), set_union2(X_60, A)) = empty_set),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(A, X_60),
                          set_union2(X_60, A)),
                        set_difference(A, set_union2(X_60, A)))), [1],
                 $fot(empty_set)]])).

cnf(refute_0_174, plain,
    (set_difference(set_union2(A, X_60), set_union2(X_60, A)) !=
     set_difference(A, set_union2(X_60, A)) |
     set_difference(set_union2(A, X_60), set_union2(X_60, A)) = empty_set),
    inference(resolve,
              [$cnf($equal(set_difference(A, set_union2(X_60, A)),
                      empty_set))], [refute_0_172, refute_0_173])).

cnf(refute_0_175, plain,
    (set_difference(set_union2(A, X_60), set_union2(X_60, A)) = empty_set),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A, X_60),
                        set_union2(X_60, A)),
                      set_difference(A, set_union2(X_60, A))))],
              [refute_0_100, refute_0_174])).

cnf(refute_0_176, plain,
    (set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) =
     empty_set),
    inference(subst, [],
              [refute_0_175 :
               [bind(A, $fot(X_188)), bind(X_60, $fot(X_187))]])).

cnf(refute_0_177, plain,
    (set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) !=
     empty_set |
     set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) !=
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) |
     empty_set =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(X_188, X_187),
                          set_union2(X_187, X_188)),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188)))), [0],
                 $fot(empty_set)]])).

cnf(refute_0_178, plain,
    (set_difference(set_union2(X_188, X_187), set_union2(X_187, X_188)) !=
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) |
     empty_set =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_union2(X_187, X_188)), empty_set))],
              [refute_0_176, refute_0_177])).

cnf(refute_0_179, plain,
    (empty_set =
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_union2(X_187, X_188)),
                      set_difference(set_difference(X_187, X_188),
                        set_union2(X_187, X_188))))],
              [refute_0_169, refute_0_178])).

cnf(refute_0_180, plain,
    (empty_set !=
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) |
     set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) = empty_set),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(empty_set)),
                bind(Y,
                     $fot(set_difference(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))))]])).

cnf(refute_0_181, plain,
    (set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) = empty_set),
    inference(resolve,
              [$cnf($equal(empty_set,
                      set_difference(set_difference(X_187, X_188),
                        set_union2(X_187, X_188))))],
              [refute_0_179, refute_0_180])).

cnf(refute_0_182, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) =
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188)))),
    introduced(tautology,
               [refl,
                [$fot(set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))))]])).

cnf(refute_0_183, plain,
    (set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) != empty_set |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) !=
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) =
     set_difference(set_union2(X_188, X_187), empty_set)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(X_188, X_187),
                          set_difference(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))),
                        set_difference(set_union2(X_188, X_187),
                          set_difference(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))))), [1, 1],
                 $fot(empty_set)]])).

cnf(refute_0_184, plain,
    (set_difference(set_difference(X_187, X_188),
       set_union2(X_187, X_188)) != empty_set |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) =
     set_difference(set_union2(X_188, X_187), empty_set)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))),
                      set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188)))))],
              [refute_0_182, refute_0_183])).

cnf(refute_0_185, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) =
     set_difference(set_union2(X_188, X_187), empty_set)),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_187, X_188),
                        set_union2(X_187, X_188)), empty_set))],
              [refute_0_181, refute_0_184])).

cnf(refute_0_186, plain,
    (set_difference(set_union2(X_188, X_187), empty_set) !=
     set_union2(X_188, X_187) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) !=
     set_difference(set_union2(X_188, X_187), empty_set) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_188, X_187)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_difference(set_union2(X_188, X_187),
                            set_difference(set_difference(X_187, X_188),
                              set_union2(X_187, X_188))))),
                bind(Y,
                     $fot(set_difference(set_union2(X_188, X_187),
                            empty_set))),
                bind(Z, $fot(set_union2(X_188, X_187)))]])).

cnf(refute_0_187, plain,
    (set_difference(set_union2(X_188, X_187), empty_set) !=
     set_union2(X_188, X_187) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_188, X_187)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))),
                      set_difference(set_union2(X_188, X_187),
                        empty_set)))], [refute_0_185, refute_0_186])).

cnf(refute_0_188, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) = set_union2(X_188, X_187)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        empty_set), set_union2(X_188, X_187)))],
              [refute_0_165, refute_0_187])).

cnf(refute_0_189, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) != set_union2(X_187, X_188) |
     set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) != set_union2(X_188, X_187) |
     set_union2(X_188, X_187) = set_union2(X_187, X_188)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(X_188, X_187),
                          set_difference(set_difference(X_187, X_188),
                            set_union2(X_187, X_188))),
                        set_union2(X_187, X_188))), [0],
                 $fot(set_union2(X_188, X_187))]])).

cnf(refute_0_190, plain,
    (set_difference(set_union2(X_188, X_187),
       set_difference(set_difference(X_187, X_188),
         set_union2(X_187, X_188))) != set_union2(X_187, X_188) |
     set_union2(X_188, X_187) = set_union2(X_187, X_188)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))),
                      set_union2(X_188, X_187)))],
              [refute_0_188, refute_0_189])).

cnf(refute_0_191, plain,
    (set_union2(X_188, X_187) = set_union2(X_187, X_188)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_188, X_187),
                        set_difference(set_difference(X_187, X_188),
                          set_union2(X_187, X_188))),
                      set_union2(X_187, X_188)))],
              [refute_0_164, refute_0_190])).

cnf(refute_0_192, plain,
    (set_union2(X_188, X_187) != set_union2(X_187, X_188) |
     set_union2(X_187, X_188) = set_union2(X_188, X_187)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(X_188, X_187))),
                bind(Y, $fot(set_union2(X_187, X_188)))]])).

cnf(refute_0_193, plain,
    (set_union2(X_187, X_188) = set_union2(X_188, X_187)),
    inference(resolve,
              [$cnf($equal(set_union2(X_188, X_187),
                      set_union2(X_187, X_188)))],
              [refute_0_191, refute_0_192])).

cnf(refute_0_194, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_union2(empty_set,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187,
                     $fot(set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4))),
                bind(X_188, $fot(empty_set))]])).

cnf(refute_0_195, plain,
    (set_union2(empty_set,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) !=
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) !=
     set_union2(empty_set,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(skolemFOFtoCNF_B_1,
                              skolemFOFtoCNF_C_4), empty_set))),
                bind(Y,
                     $fot(set_union2(empty_set,
                            set_difference(skolemFOFtoCNF_B_1,
                              skolemFOFtoCNF_C_4)))),
                bind(Z,
                     $fot(set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4)))]])).

cnf(refute_0_196, plain,
    (set_union2(empty_set,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) !=
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4), empty_set),
                      set_union2(empty_set,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4))))],
              [refute_0_194, refute_0_195])).

cnf(refute_0_197, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) =
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4)),
                      set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4)))],
              [refute_0_37, refute_0_196])).

cnf(refute_0_198, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) !=
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) |
     set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4), empty_set),
                        set_union2(skolemFOFtoCNF_B_1,
                          set_difference(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_C_4)))), [0],
                 $fot(set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4))]])).

cnf(refute_0_199, plain,
    (set_union2(set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4),
       empty_set) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4), empty_set),
                      set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4)))],
              [refute_0_197, refute_0_198])).

cnf(refute_0_200, plain,
    (set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4), empty_set),
                      set_union2(skolemFOFtoCNF_B_1,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4))))],
              [refute_0_31, refute_0_199])).

cnf(refute_0_201, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       set_difference(set_difference(X_63, X_237), X_63)) = X_63),
    inference(subst, [],
              [refute_0_121 :
               [bind(A, $fot(set_difference(X_63, X_237)))]])).

cnf(refute_0_202, plain, (subset(set_difference(A, B), A)),
    inference(canonicalize, [], [normalize_0_30])).

cnf(refute_0_203, plain, (subset(set_difference(X_236, B), X_236)),
    inference(subst, [], [refute_0_202 : [bind(A, $fot(X_236))]])).

cnf(refute_0_204, plain,
    (~ subset(A, B) | set_difference(A, B) = empty_set),
    inference(canonicalize, [], [normalize_0_34])).

cnf(refute_0_205, plain,
    (~ subset(set_difference(X_236, B), X_236) |
     set_difference(set_difference(X_236, B), X_236) = empty_set),
    inference(subst, [],
              [refute_0_204 :
               [bind(A, $fot(set_difference(X_236, B))),
                bind(B, $fot(X_236))]])).

cnf(refute_0_206, plain,
    (set_difference(set_difference(X_236, B), X_236) = empty_set),
    inference(resolve, [$cnf(subset(set_difference(X_236, B), X_236))],
              [refute_0_203, refute_0_205])).

cnf(refute_0_207, plain,
    (set_difference(set_difference(X_63, X_237), X_63) = empty_set),
    inference(subst, [],
              [refute_0_206 :
               [bind(B, $fot(X_237)), bind(X_236, $fot(X_63))]])).

cnf(refute_0_208, plain,
    (set_difference(set_difference(X_63, X_237), X_63) != empty_set |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       set_difference(set_difference(X_63, X_237), X_63)) != X_63 |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = X_63),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                              X_237), X_63),
                          set_difference(set_difference(X_63, X_237),
                            X_63)), X_63)), [0, 1], $fot(empty_set)]])).

cnf(refute_0_209, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       set_difference(set_difference(X_63, X_237), X_63)) != X_63 |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = X_63),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_63, X_237),
                        X_63), empty_set))],
              [refute_0_207, refute_0_208])).

cnf(refute_0_210, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = X_63),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                            X_237), X_63),
                        set_difference(set_difference(X_63, X_237), X_63)),
                      X_63))], [refute_0_201, refute_0_209])).

cnf(refute_0_211, plain,
    (set_union2(set_difference(X_63, X_237), X_63) =
     set_union2(X_63, set_difference(X_63, X_237))),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187, $fot(set_difference(X_63, X_237))),
                bind(X_188, $fot(X_63))]])).

cnf(refute_0_212, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = set_union2(set_difference(X_63, X_237), X_63)),
    inference(subst, [],
              [refute_0_112 :
               [bind(A,
                     $fot(set_union2(set_difference(X_63, X_237),
                            X_63)))]])).

cnf(refute_0_213, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) != set_union2(set_difference(X_63, X_237), X_63) |
     set_union2(set_difference(X_63, X_237), X_63) !=
     set_union2(X_63, set_difference(X_63, X_237)) |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = set_union2(X_63, set_difference(X_63, X_237))),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_difference(set_union2(set_difference(X_63,
                                X_237), X_63), empty_set))),
                bind(Y,
                     $fot(set_union2(set_difference(X_63, X_237), X_63))),
                bind(Z,
                     $fot(set_union2(X_63,
                            set_difference(X_63, X_237))))]])).

cnf(refute_0_214, plain,
    (set_union2(set_difference(X_63, X_237), X_63) !=
     set_union2(X_63, set_difference(X_63, X_237)) |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = set_union2(X_63, set_difference(X_63, X_237))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                            X_237), X_63), empty_set),
                      set_union2(set_difference(X_63, X_237), X_63)))],
              [refute_0_212, refute_0_213])).

cnf(refute_0_215, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) = set_union2(X_63, set_difference(X_63, X_237))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_63, X_237), X_63),
                      set_union2(X_63, set_difference(X_63, X_237))))],
              [refute_0_211, refute_0_214])).

cnf(refute_0_216, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) != X_63 |
     set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) != set_union2(X_63, set_difference(X_63, X_237)) |
     set_union2(X_63, set_difference(X_63, X_237)) = X_63),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                              X_237), X_63), empty_set), X_63)), [0],
                 $fot(set_union2(X_63, set_difference(X_63, X_237)))]])).

cnf(refute_0_217, plain,
    (set_difference(set_union2(set_difference(X_63, X_237), X_63),
       empty_set) != X_63 |
     set_union2(X_63, set_difference(X_63, X_237)) = X_63),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                            X_237), X_63), empty_set),
                      set_union2(X_63, set_difference(X_63, X_237))))],
              [refute_0_215, refute_0_216])).

cnf(refute_0_218, plain,
    (set_union2(X_63, set_difference(X_63, X_237)) = X_63),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(set_difference(X_63,
                            X_237), X_63), empty_set), X_63))],
              [refute_0_210, refute_0_217])).

cnf(refute_0_219, plain,
    (set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) =
     skolemFOFtoCNF_B_1),
    inference(subst, [],
              [refute_0_218 :
               [bind(X_237, $fot(skolemFOFtoCNF_C_4)),
                bind(X_63, $fot(skolemFOFtoCNF_B_1))]])).

cnf(refute_0_220, plain,
    (set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) !=
     skolemFOFtoCNF_B_1 |
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     skolemFOFtoCNF_B_1),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4), skolemFOFtoCNF_B_1)), [0],
                 $fot(set_union2(skolemFOFtoCNF_B_1,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4)))]])).

cnf(refute_0_221, plain,
    (set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) !=
     set_union2(skolemFOFtoCNF_B_1,
       set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4)) |
     set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     skolemFOFtoCNF_B_1),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_B_1,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4)), skolemFOFtoCNF_B_1))],
              [refute_0_219, refute_0_220])).

cnf(refute_0_222, plain,
    (set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) =
     skolemFOFtoCNF_B_1),
    inference(resolve,
              [$cnf($equal(set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4),
                      set_union2(skolemFOFtoCNF_B_1,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4))))],
              [refute_0_200, refute_0_221])).

cnf(refute_0_223, plain,
    (set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4) !=
     skolemFOFtoCNF_B_1 | ~ in(X_284, skolemFOFtoCNF_B_1) |
     in(X_284, set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    introduced(tautology,
               [equality,
                [$cnf(~ in(X_284,
                        set_difference(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_C_4))), [1],
                 $fot(skolemFOFtoCNF_B_1)]])).

cnf(refute_0_224, plain,
    (~ in(X_284, skolemFOFtoCNF_B_1) |
     in(X_284, set_difference(skolemFOFtoCNF_B_1, skolemFOFtoCNF_C_4))),
    inference(resolve,
              [$cnf($equal(set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4), skolemFOFtoCNF_B_1))],
              [refute_0_222, refute_0_223])).

cnf(refute_0_225, plain,
    (~ in(X_284, skolemFOFtoCNF_B_1) | ~ in(X_284, skolemFOFtoCNF_C_4)),
    inference(resolve,
              [$cnf(in(X_284,
                      set_difference(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_C_4)))],
              [refute_0_224, refute_0_6])).

cnf(refute_0_226, plain,
    (~
       in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, X_328),
         skolemFOFtoCNF_B_1) |
     ~
       in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, X_328),
         skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_225 :
               [bind(X_284,
                     $fot(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4,
                            X_328)))]])).

cnf(refute_0_227, plain,
    (disjoint(A, B) |
     in(skolemFOFtoCNF_C_3(A, B), set_intersection2(A, B))),
    inference(canonicalize, [], [normalize_0_38])).

cnf(refute_0_228, plain,
    (disjoint(X_324, X_325) |
     in(skolemFOFtoCNF_C_3(X_324, X_325),
       set_intersection2(X_324, X_325))),
    inference(subst, [],
              [refute_0_227 :
               [bind(A, $fot(X_324)), bind(B, $fot(X_325))]])).

cnf(refute_0_229, plain,
    (C != set_intersection2(A, B) | ~ in(D, C) | in(D, A)),
    inference(canonicalize, [], [normalize_0_42])).

cnf(refute_0_230, plain,
    (set_intersection2(A, B) != set_intersection2(A, B) |
     ~ in(D, set_intersection2(A, B)) | in(D, A)),
    inference(subst, [],
              [refute_0_229 : [bind(C, $fot(set_intersection2(A, B)))]])).

cnf(refute_0_231, plain,
    (set_intersection2(A, B) = set_intersection2(A, B)),
    introduced(tautology, [refl, [$fot(set_intersection2(A, B))]])).

cnf(refute_0_232, plain, (~ in(D, set_intersection2(A, B)) | in(D, A)),
    inference(resolve,
              [$cnf($equal(set_intersection2(A, B),
                      set_intersection2(A, B)))],
              [refute_0_231, refute_0_230])).

cnf(refute_0_233, plain,
    (~
       in(skolemFOFtoCNF_C_3(X_324, X_325),
         set_intersection2(X_324, X_325)) |
     in(skolemFOFtoCNF_C_3(X_324, X_325), X_324)),
    inference(subst, [],
              [refute_0_232 :
               [bind(A, $fot(X_324)), bind(B, $fot(X_325)),
                bind(D, $fot(skolemFOFtoCNF_C_3(X_324, X_325)))]])).

cnf(refute_0_234, plain,
    (disjoint(X_324, X_325) | in(skolemFOFtoCNF_C_3(X_324, X_325), X_324)),
    inference(resolve,
              [$cnf(in(skolemFOFtoCNF_C_3(X_324, X_325),
                      set_intersection2(X_324, X_325)))],
              [refute_0_228, refute_0_233])).

cnf(refute_0_235, plain,
    (disjoint(skolemFOFtoCNF_C_4, X_328) |
     in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, X_328),
       skolemFOFtoCNF_C_4)),
    inference(subst, [],
              [refute_0_234 :
               [bind(X_324, $fot(skolemFOFtoCNF_C_4)),
                bind(X_325, $fot(X_328))]])).

cnf(refute_0_236, plain,
    (~
       in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, X_328),
         skolemFOFtoCNF_B_1) | disjoint(skolemFOFtoCNF_C_4, X_328)),
    inference(resolve,
              [$cnf(in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, X_328),
                      skolemFOFtoCNF_C_4))],
              [refute_0_235, refute_0_226])).

cnf(refute_0_237, plain,
    (~
       in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2),
         skolemFOFtoCNF_B_1) |
     disjoint(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2)),
    inference(subst, [],
              [refute_0_236 : [bind(X_328, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_238, plain,
    (C != set_intersection2(A, B) | ~ in(D, C) | in(D, B)),
    inference(canonicalize, [], [normalize_0_43])).

cnf(refute_0_239, plain,
    (set_intersection2(A, B) != set_intersection2(A, B) |
     ~ in(D, set_intersection2(A, B)) | in(D, B)),
    inference(subst, [],
              [refute_0_238 : [bind(C, $fot(set_intersection2(A, B)))]])).

cnf(refute_0_240, plain, (~ in(D, set_intersection2(A, B)) | in(D, B)),
    inference(resolve,
              [$cnf($equal(set_intersection2(A, B),
                      set_intersection2(A, B)))],
              [refute_0_231, refute_0_239])).

cnf(refute_0_241, plain,
    (~
       in(X_336,
         set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) |
     in(X_336, skolemFOFtoCNF_B_1)),
    inference(subst, [],
              [refute_0_240 :
               [bind(A, $fot(skolemFOFtoCNF_A_2)),
                bind(B, $fot(skolemFOFtoCNF_B_1)),
                bind(D, $fot(X_336))]])).

cnf(refute_0_242, plain,
    (set_union2(set_difference(X_114, X_113),
       set_difference(set_intersection2(X_113, X_114),
         set_difference(X_114, X_113))) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    inference(subst, [],
              [refute_0_7 :
               [bind(A, $fot(set_difference(X_114, X_113))),
                bind(B, $fot(set_intersection2(X_113, X_114)))]])).

cnf(refute_0_243, plain,
    (set_difference(set_union2(B, X_87),
       set_difference(set_union2(B, X_87), B)) =
     set_intersection2(set_union2(B, X_87), B)),
    inference(subst, [],
              [refute_0_9 : [bind(A, $fot(set_union2(B, X_87)))]])).

cnf(refute_0_244, plain,
    (set_difference(set_union2(X_50, X_51), X_51) =
     set_difference(X_50, X_51)),
    inference(subst, [],
              [refute_0_39 : [bind(A, $fot(X_50)), bind(B, $fot(X_51))]])).

cnf(refute_0_245, plain, (set_union2(X_51, X_50) = set_union2(X_50, X_51)),
    inference(subst, [],
              [refute_0_14 : [bind(A, $fot(X_51)), bind(B, $fot(X_50))]])).

cnf(refute_0_246, plain,
    (set_union2(X_51, X_50) != set_union2(X_50, X_51) |
     set_union2(X_50, X_51) = set_union2(X_51, X_50)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(X_51, X_50))),
                bind(Y, $fot(set_union2(X_50, X_51)))]])).

cnf(refute_0_247, plain, (set_union2(X_50, X_51) = set_union2(X_51, X_50)),
    inference(resolve,
              [$cnf($equal(set_union2(X_51, X_50),
                      set_union2(X_50, X_51)))],
              [refute_0_245, refute_0_246])).

cnf(refute_0_248, plain,
    (set_difference(set_union2(X_50, X_51), X_51) !=
     set_difference(X_50, X_51) |
     set_union2(X_50, X_51) != set_union2(X_51, X_50) |
     set_difference(set_union2(X_51, X_50), X_51) =
     set_difference(X_50, X_51)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(X_50, X_51), X_51),
                        set_difference(X_50, X_51))), [0, 0],
                 $fot(set_union2(X_51, X_50))]])).

cnf(refute_0_249, plain,
    (set_difference(set_union2(X_50, X_51), X_51) !=
     set_difference(X_50, X_51) |
     set_difference(set_union2(X_51, X_50), X_51) =
     set_difference(X_50, X_51)),
    inference(resolve,
              [$cnf($equal(set_union2(X_50, X_51),
                      set_union2(X_51, X_50)))],
              [refute_0_247, refute_0_248])).

cnf(refute_0_250, plain,
    (set_difference(set_union2(X_51, X_50), X_51) =
     set_difference(X_50, X_51)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(X_50, X_51), X_51),
                      set_difference(X_50, X_51)))],
              [refute_0_244, refute_0_249])).

cnf(refute_0_251, plain,
    (set_difference(set_union2(B, X_87), B) = set_difference(X_87, B)),
    inference(subst, [],
              [refute_0_250 :
               [bind(X_50, $fot(X_87)), bind(X_51, $fot(B))]])).

cnf(refute_0_252, plain,
    (set_difference(set_union2(B, X_87), B) != set_difference(X_87, B) |
     set_difference(set_union2(B, X_87),
       set_difference(set_union2(B, X_87), B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(set_union2(B, X_87), set_difference(X_87, B)) =
     set_intersection2(set_union2(B, X_87), B)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(B, X_87),
                          set_difference(set_union2(B, X_87), B)),
                        set_intersection2(set_union2(B, X_87), B))),
                 [0, 1], $fot(set_difference(X_87, B))]])).

cnf(refute_0_253, plain,
    (set_difference(set_union2(B, X_87),
       set_difference(set_union2(B, X_87), B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(set_union2(B, X_87), set_difference(X_87, B)) =
     set_intersection2(set_union2(B, X_87), B)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(B, X_87), B),
                      set_difference(X_87, B)))],
              [refute_0_251, refute_0_252])).

cnf(refute_0_254, plain,
    (set_difference(set_union2(B, X_87), set_difference(X_87, B)) =
     set_intersection2(set_union2(B, X_87), B)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(B, X_87),
                        set_difference(set_union2(B, X_87), B)),
                      set_intersection2(set_union2(B, X_87), B)))],
              [refute_0_243, refute_0_253])).

cnf(refute_0_255, plain,
    (set_difference(set_union2(A, set_difference(X_59, A)),
       set_difference(X_59, A)) =
     set_difference(A, set_difference(X_59, A))),
    inference(subst, [],
              [refute_0_39 : [bind(B, $fot(set_difference(X_59, A)))]])).

cnf(refute_0_256, plain,
    (set_union2(A, set_difference(X_59, A)) = set_union2(A, X_59)),
    inference(subst, [], [refute_0_7 : [bind(B, $fot(X_59))]])).

cnf(refute_0_257, plain,
    (set_difference(set_union2(A, set_difference(X_59, A)),
       set_difference(X_59, A)) !=
     set_difference(A, set_difference(X_59, A)) |
     set_union2(A, set_difference(X_59, A)) != set_union2(A, X_59) |
     set_difference(set_union2(A, X_59), set_difference(X_59, A)) =
     set_difference(A, set_difference(X_59, A))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(A,
                            set_difference(X_59, A)),
                          set_difference(X_59, A)),
                        set_difference(A, set_difference(X_59, A)))),
                 [0, 0], $fot(set_union2(A, X_59))]])).

cnf(refute_0_258, plain,
    (set_difference(set_union2(A, set_difference(X_59, A)),
       set_difference(X_59, A)) !=
     set_difference(A, set_difference(X_59, A)) |
     set_difference(set_union2(A, X_59), set_difference(X_59, A)) =
     set_difference(A, set_difference(X_59, A))),
    inference(resolve,
              [$cnf($equal(set_union2(A, set_difference(X_59, A)),
                      set_union2(A, X_59)))],
              [refute_0_256, refute_0_257])).

cnf(refute_0_259, plain,
    (set_difference(set_union2(A, X_59), set_difference(X_59, A)) =
     set_difference(A, set_difference(X_59, A))),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(A,
                          set_difference(X_59, A)),
                        set_difference(X_59, A)),
                      set_difference(A, set_difference(X_59, A))))],
              [refute_0_255, refute_0_258])).

cnf(refute_0_260, plain,
    (set_difference(set_union2(B, X_87), set_difference(X_87, B)) =
     set_difference(B, set_difference(X_87, B))),
    inference(subst, [],
              [refute_0_259 :
               [bind(A, $fot(B)), bind(X_59, $fot(X_87))]])).

cnf(refute_0_261, plain,
    (set_difference(set_union2(B, X_87), set_difference(X_87, B)) !=
     set_difference(B, set_difference(X_87, B)) |
     set_difference(set_union2(B, X_87), set_difference(X_87, B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(B, set_difference(X_87, B)) =
     set_intersection2(set_union2(B, X_87), B)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_union2(B, X_87),
                          set_difference(X_87, B)),
                        set_intersection2(set_union2(B, X_87), B))), [0],
                 $fot(set_difference(B, set_difference(X_87, B)))]])).

cnf(refute_0_262, plain,
    (set_difference(set_union2(B, X_87), set_difference(X_87, B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(B, set_difference(X_87, B)) =
     set_intersection2(set_union2(B, X_87), B)),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(B, X_87),
                        set_difference(X_87, B)),
                      set_difference(B, set_difference(X_87, B))))],
              [refute_0_260, refute_0_261])).

cnf(refute_0_263, plain,
    (X_73 = set_intersection2(X_73, set_union2(X_74, X_73))),
    inference(subst, [],
              [refute_0_115 :
               [bind(A, $fot(X_73)), bind(X_71, $fot(X_74))]])).

cnf(refute_0_264, plain, (set_union2(X_73, X_74) = set_union2(X_74, X_73)),
    inference(subst, [],
              [refute_0_14 : [bind(A, $fot(X_73)), bind(B, $fot(X_74))]])).

cnf(refute_0_265, plain,
    (set_union2(X_73, X_74) != set_union2(X_74, X_73) |
     set_union2(X_74, X_73) = set_union2(X_73, X_74)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_union2(X_73, X_74))),
                bind(Y, $fot(set_union2(X_74, X_73)))]])).

cnf(refute_0_266, plain, (set_union2(X_74, X_73) = set_union2(X_73, X_74)),
    inference(resolve,
              [$cnf($equal(set_union2(X_73, X_74),
                      set_union2(X_74, X_73)))],
              [refute_0_264, refute_0_265])).

cnf(refute_0_267, plain,
    (X_73 != set_intersection2(X_73, set_union2(X_74, X_73)) |
     set_union2(X_74, X_73) != set_union2(X_73, X_74) |
     X_73 = set_intersection2(X_73, set_union2(X_73, X_74))),
    introduced(tautology,
               [equality,
                [$cnf($equal(X_73,
                        set_intersection2(X_73, set_union2(X_74, X_73)))),
                 [1, 1], $fot(set_union2(X_73, X_74))]])).

cnf(refute_0_268, plain,
    (X_73 != set_intersection2(X_73, set_union2(X_74, X_73)) |
     X_73 = set_intersection2(X_73, set_union2(X_73, X_74))),
    inference(resolve,
              [$cnf($equal(set_union2(X_74, X_73),
                      set_union2(X_73, X_74)))],
              [refute_0_266, refute_0_267])).

cnf(refute_0_269, plain,
    (X_73 = set_intersection2(X_73, set_union2(X_73, X_74))),
    inference(resolve,
              [$cnf($equal(X_73,
                      set_intersection2(X_73, set_union2(X_74, X_73))))],
              [refute_0_263, refute_0_268])).

cnf(refute_0_270, plain,
    (X_73 != set_intersection2(X_73, set_union2(X_73, X_74)) |
     set_intersection2(X_73, set_union2(X_73, X_74)) = X_73),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(X_73)),
                bind(Y,
                     $fot(set_intersection2(X_73,
                            set_union2(X_73, X_74))))]])).

cnf(refute_0_271, plain,
    (set_intersection2(X_73, set_union2(X_73, X_74)) = X_73),
    inference(resolve,
              [$cnf($equal(X_73,
                      set_intersection2(X_73, set_union2(X_73, X_74))))],
              [refute_0_269, refute_0_270])).

cnf(refute_0_272, plain, (set_intersection2(B, set_union2(B, X_87)) = B),
    inference(subst, [],
              [refute_0_271 :
               [bind(X_73, $fot(B)), bind(X_74, $fot(X_87))]])).

cnf(refute_0_273, plain,
    (set_intersection2(set_union2(B, X_87), B) =
     set_intersection2(B, set_union2(B, X_87))),
    inference(subst, [],
              [refute_0_46 :
               [bind(A, $fot(B)), bind(B, $fot(set_union2(B, X_87)))]])).

cnf(refute_0_274, plain,
    (set_intersection2(B, set_union2(B, X_87)) != B |
     set_intersection2(set_union2(B, X_87), B) !=
     set_intersection2(B, set_union2(B, X_87)) |
     set_intersection2(set_union2(B, X_87), B) = B),
    inference(subst, [],
              [refute_0_71 :
               [bind(X, $fot(set_intersection2(set_union2(B, X_87), B))),
                bind(Y, $fot(set_intersection2(B, set_union2(B, X_87)))),
                bind(Z, $fot(B))]])).

cnf(refute_0_275, plain,
    (set_intersection2(B, set_union2(B, X_87)) != B |
     set_intersection2(set_union2(B, X_87), B) = B),
    inference(resolve,
              [$cnf($equal(set_intersection2(set_union2(B, X_87), B),
                      set_intersection2(B, set_union2(B, X_87))))],
              [refute_0_273, refute_0_274])).

cnf(refute_0_276, plain, (set_intersection2(set_union2(B, X_87), B) = B),
    inference(resolve,
              [$cnf($equal(set_intersection2(B, set_union2(B, X_87)), B))],
              [refute_0_272, refute_0_275])).

cnf(refute_0_277, plain,
    (set_difference(B, set_difference(X_87, B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_intersection2(set_union2(B, X_87), B) != B |
     set_difference(B, set_difference(X_87, B)) = B),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(B, set_difference(X_87, B)),
                        B)), [0],
                 $fot(set_intersection2(set_union2(B, X_87), B))]])).

cnf(refute_0_278, plain,
    (set_difference(B, set_difference(X_87, B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(B, set_difference(X_87, B)) = B),
    inference(resolve,
              [$cnf($equal(set_intersection2(set_union2(B, X_87), B), B))],
              [refute_0_276, refute_0_277])).

cnf(refute_0_279, plain,
    (set_difference(set_union2(B, X_87), set_difference(X_87, B)) !=
     set_intersection2(set_union2(B, X_87), B) |
     set_difference(B, set_difference(X_87, B)) = B),
    inference(resolve,
              [$cnf($equal(set_difference(B, set_difference(X_87, B)),
                      set_intersection2(set_union2(B, X_87), B)))],
              [refute_0_262, refute_0_278])).

cnf(refute_0_280, plain, (set_difference(B, set_difference(X_87, B)) = B),
    inference(resolve,
              [$cnf($equal(set_difference(set_union2(B, X_87),
                        set_difference(X_87, B)),
                      set_intersection2(set_union2(B, X_87), B)))],
              [refute_0_254, refute_0_279])).

cnf(refute_0_281, plain,
    (set_difference(set_intersection2(X_102, X_101),
       set_difference(set_difference(X_102, X_101),
         set_intersection2(X_102, X_101))) =
     set_intersection2(X_102, X_101)),
    inference(subst, [],
              [refute_0_280 :
               [bind(B, $fot(set_intersection2(X_102, X_101))),
                bind(X_87, $fot(set_difference(X_102, X_101)))]])).

cnf(refute_0_282, plain,
    (set_difference(set_difference(X_90, B),
       set_difference(X_90, set_difference(X_90, B))) =
     set_difference(X_90, B)),
    inference(subst, [],
              [refute_0_280 :
               [bind(B, $fot(set_difference(X_90, B))),
                bind(X_87, $fot(X_90))]])).

cnf(refute_0_283, plain,
    (set_difference(X_90, set_difference(X_90, B)) =
     set_intersection2(X_90, B)),
    inference(subst, [], [refute_0_9 : [bind(A, $fot(X_90))]])).

cnf(refute_0_284, plain,
    (set_difference(X_90, set_difference(X_90, B)) !=
     set_intersection2(X_90, B) |
     set_difference(set_difference(X_90, B),
       set_difference(X_90, set_difference(X_90, B))) !=
     set_difference(X_90, B) |
     set_difference(set_difference(X_90, B), set_intersection2(X_90, B)) =
     set_difference(X_90, B)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_difference(X_90, B),
                          set_difference(X_90, set_difference(X_90, B))),
                        set_difference(X_90, B))), [0, 1],
                 $fot(set_intersection2(X_90, B))]])).

cnf(refute_0_285, plain,
    (set_difference(set_difference(X_90, B),
       set_difference(X_90, set_difference(X_90, B))) !=
     set_difference(X_90, B) |
     set_difference(set_difference(X_90, B), set_intersection2(X_90, B)) =
     set_difference(X_90, B)),
    inference(resolve,
              [$cnf($equal(set_difference(X_90, set_difference(X_90, B)),
                      set_intersection2(X_90, B)))],
              [refute_0_283, refute_0_284])).

cnf(refute_0_286, plain,
    (set_difference(set_difference(X_90, B), set_intersection2(X_90, B)) =
     set_difference(X_90, B)),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_90, B),
                        set_difference(X_90, set_difference(X_90, B))),
                      set_difference(X_90, B)))],
              [refute_0_282, refute_0_285])).

cnf(refute_0_287, plain,
    (set_difference(set_difference(X_102, X_101),
       set_intersection2(X_102, X_101)) = set_difference(X_102, X_101)),
    inference(subst, [],
              [refute_0_286 :
               [bind(B, $fot(X_101)), bind(X_90, $fot(X_102))]])).

cnf(refute_0_288, plain,
    (set_difference(set_difference(X_102, X_101),
       set_intersection2(X_102, X_101)) != set_difference(X_102, X_101) |
     set_difference(set_intersection2(X_102, X_101),
       set_difference(set_difference(X_102, X_101),
         set_intersection2(X_102, X_101))) !=
     set_intersection2(X_102, X_101) |
     set_difference(set_intersection2(X_102, X_101),
       set_difference(X_102, X_101)) = set_intersection2(X_102, X_101)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_intersection2(X_102,
                            X_101),
                          set_difference(set_difference(X_102, X_101),
                            set_intersection2(X_102, X_101))),
                        set_intersection2(X_102, X_101))), [0, 1],
                 $fot(set_difference(X_102, X_101))]])).

cnf(refute_0_289, plain,
    (set_difference(set_intersection2(X_102, X_101),
       set_difference(set_difference(X_102, X_101),
         set_intersection2(X_102, X_101))) !=
     set_intersection2(X_102, X_101) |
     set_difference(set_intersection2(X_102, X_101),
       set_difference(X_102, X_101)) = set_intersection2(X_102, X_101)),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_102, X_101),
                        set_intersection2(X_102, X_101)),
                      set_difference(X_102, X_101)))],
              [refute_0_287, refute_0_288])).

cnf(refute_0_290, plain,
    (set_difference(set_intersection2(X_102, X_101),
       set_difference(X_102, X_101)) = set_intersection2(X_102, X_101)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_102, X_101),
                        set_difference(set_difference(X_102, X_101),
                          set_intersection2(X_102, X_101))),
                      set_intersection2(X_102, X_101)))],
              [refute_0_281, refute_0_289])).

cnf(refute_0_291, plain,
    (set_difference(set_intersection2(X_108, X_107),
       set_difference(X_108, X_107)) = set_intersection2(X_108, X_107)),
    inference(subst, [],
              [refute_0_290 :
               [bind(X_101, $fot(X_107)), bind(X_102, $fot(X_108))]])).

cnf(refute_0_292, plain,
    (set_intersection2(X_107, X_108) = set_intersection2(X_108, X_107)),
    inference(subst, [],
              [refute_0_44 :
               [bind(A, $fot(X_107)), bind(B, $fot(X_108))]])).

cnf(refute_0_293, plain,
    (set_intersection2(X_107, X_108) != set_intersection2(X_108, X_107) |
     set_intersection2(X_108, X_107) = set_intersection2(X_107, X_108)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_intersection2(X_107, X_108))),
                bind(Y, $fot(set_intersection2(X_108, X_107)))]])).

cnf(refute_0_294, plain,
    (set_intersection2(X_108, X_107) = set_intersection2(X_107, X_108)),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_107, X_108),
                      set_intersection2(X_108, X_107)))],
              [refute_0_292, refute_0_293])).

cnf(refute_0_295, plain,
    (set_difference(set_intersection2(X_108, X_107),
       set_difference(X_108, X_107)) != set_intersection2(X_108, X_107) |
     set_intersection2(X_108, X_107) != set_intersection2(X_107, X_108) |
     set_difference(set_intersection2(X_107, X_108),
       set_difference(X_108, X_107)) = set_intersection2(X_108, X_107)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_intersection2(X_108,
                            X_107), set_difference(X_108, X_107)),
                        set_intersection2(X_108, X_107))), [0, 0],
                 $fot(set_intersection2(X_107, X_108))]])).

cnf(refute_0_296, plain,
    (set_difference(set_intersection2(X_108, X_107),
       set_difference(X_108, X_107)) != set_intersection2(X_108, X_107) |
     set_difference(set_intersection2(X_107, X_108),
       set_difference(X_108, X_107)) = set_intersection2(X_108, X_107)),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_108, X_107),
                      set_intersection2(X_107, X_108)))],
              [refute_0_294, refute_0_295])).

cnf(refute_0_297, plain,
    (set_difference(set_intersection2(X_107, X_108),
       set_difference(X_108, X_107)) = set_intersection2(X_108, X_107)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_108, X_107),
                        set_difference(X_108, X_107)),
                      set_intersection2(X_108, X_107)))],
              [refute_0_291, refute_0_296])).

cnf(refute_0_298, plain,
    (set_difference(set_intersection2(X_113, X_114),
       set_difference(X_114, X_113)) = set_intersection2(X_114, X_113)),
    inference(subst, [],
              [refute_0_297 :
               [bind(X_107, $fot(X_113)), bind(X_108, $fot(X_114))]])).

cnf(refute_0_299, plain,
    (set_difference(set_intersection2(X_113, X_114),
       set_difference(X_114, X_113)) != set_intersection2(X_114, X_113) |
     set_union2(set_difference(X_114, X_113),
       set_difference(set_intersection2(X_113, X_114),
         set_difference(X_114, X_113))) !=
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114)) |
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_114, X_113),
                          set_difference(set_intersection2(X_113, X_114),
                            set_difference(X_114, X_113))),
                        set_union2(set_difference(X_114, X_113),
                          set_intersection2(X_113, X_114)))), [0, 1],
                 $fot(set_intersection2(X_114, X_113))]])).

cnf(refute_0_300, plain,
    (set_union2(set_difference(X_114, X_113),
       set_difference(set_intersection2(X_113, X_114),
         set_difference(X_114, X_113))) !=
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114)) |
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_113, X_114),
                        set_difference(X_114, X_113)),
                      set_intersection2(X_114, X_113)))],
              [refute_0_298, refute_0_299])).

cnf(refute_0_301, plain,
    (set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_114, X_113),
                        set_difference(set_intersection2(X_113, X_114),
                          set_difference(X_114, X_113))),
                      set_union2(set_difference(X_114, X_113),
                        set_intersection2(X_113, X_114))))],
              [refute_0_242, refute_0_300])).

cnf(refute_0_302, plain,
    (set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) =
     set_union2(X_114, set_difference(X_114, X_113))),
    inference(subst, [],
              [refute_0_23 :
               [bind(B, $fot(X_114)), bind(X_63, $fot(X_113))]])).

cnf(refute_0_303, plain,
    (set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) !=
     set_union2(X_114, set_difference(X_114, X_113)) |
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) !=
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114)) |
     set_union2(X_114, set_difference(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_114, X_113),
                          set_intersection2(X_114, X_113)),
                        set_union2(set_difference(X_114, X_113),
                          set_intersection2(X_113, X_114)))), [0],
                 $fot(set_union2(X_114, set_difference(X_114, X_113)))]])).

cnf(refute_0_304, plain,
    (set_union2(set_difference(X_114, X_113),
       set_intersection2(X_114, X_113)) !=
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114)) |
     set_union2(X_114, set_difference(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_114, X_113),
                        set_intersection2(X_114, X_113)),
                      set_union2(X_114, set_difference(X_114, X_113))))],
              [refute_0_302, refute_0_303])).

cnf(refute_0_305, plain,
    (set_union2(X_114, set_difference(X_114, X_113)) =
     set_union2(set_difference(X_114, X_113),
       set_intersection2(X_113, X_114))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_114, X_113),
                        set_intersection2(X_114, X_113)),
                      set_union2(set_difference(X_114, X_113),
                        set_intersection2(X_113, X_114))))],
              [refute_0_301, refute_0_304])).

cnf(refute_0_306, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    inference(subst, [],
              [refute_0_305 :
               [bind(X_113, $fot(skolemFOFtoCNF_B_1)),
                bind(X_114, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_307, plain, (subset(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(canonicalize, [], [normalize_0_44])).

cnf(refute_0_308, plain,
    (~ subset(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) |
     set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) = empty_set),
    inference(subst, [],
              [refute_0_204 :
               [bind(A, $fot(skolemFOFtoCNF_A_2)),
                bind(B, $fot(skolemFOFtoCNF_B_1))]])).

cnf(refute_0_309, plain,
    (set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) = empty_set),
    inference(resolve,
              [$cnf(subset(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1))],
              [refute_0_307, refute_0_308])).

cnf(refute_0_310, plain,
    (set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) != empty_set |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                          set_difference(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)),
                        set_union2(set_difference(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1),
                          set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_A_2)))), [1, 0],
                 $fot(empty_set)]])).

cnf(refute_0_311, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1),
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    inference(resolve,
              [$cnf($equal(set_difference(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1), empty_set))],
              [refute_0_309, refute_0_310])).

cnf(refute_0_312, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)),
                      set_union2(set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1),
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_A_2))))],
              [refute_0_306, refute_0_311])).

cnf(refute_0_313, plain,
    (set_union2(empty_set, skolemFOFtoCNF_A_2) = skolemFOFtoCNF_A_2),
    inference(subst, [],
              [refute_0_36 : [bind(A, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_314, plain,
    (set_union2(skolemFOFtoCNF_A_2, empty_set) =
     set_union2(empty_set, skolemFOFtoCNF_A_2)),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187, $fot(skolemFOFtoCNF_A_2)),
                bind(X_188, $fot(empty_set))]])).

cnf(refute_0_315, plain,
    (set_union2(empty_set, skolemFOFtoCNF_A_2) != skolemFOFtoCNF_A_2 |
     set_union2(skolemFOFtoCNF_A_2, empty_set) !=
     set_union2(empty_set, skolemFOFtoCNF_A_2) |
     set_union2(skolemFOFtoCNF_A_2, empty_set) = skolemFOFtoCNF_A_2),
    inference(subst, [],
              [refute_0_71 :
               [bind(X, $fot(set_union2(skolemFOFtoCNF_A_2, empty_set))),
                bind(Y, $fot(set_union2(empty_set, skolemFOFtoCNF_A_2))),
                bind(Z, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_316, plain,
    (set_union2(empty_set, skolemFOFtoCNF_A_2) != skolemFOFtoCNF_A_2 |
     set_union2(skolemFOFtoCNF_A_2, empty_set) = skolemFOFtoCNF_A_2),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2, empty_set),
                      set_union2(empty_set, skolemFOFtoCNF_A_2)))],
              [refute_0_314, refute_0_315])).

cnf(refute_0_317, plain,
    (set_union2(skolemFOFtoCNF_A_2, empty_set) = skolemFOFtoCNF_A_2),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set, skolemFOFtoCNF_A_2),
                      skolemFOFtoCNF_A_2))],
              [refute_0_313, refute_0_316])).

cnf(refute_0_318, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1))),
    introduced(tautology,
               [refl,
                [$fot(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)))]])).

cnf(refute_0_319, plain,
    (set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) != empty_set |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(skolemFOFtoCNF_A_2, empty_set)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                          set_difference(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)),
                        set_union2(skolemFOFtoCNF_A_2,
                          set_difference(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)))), [1, 1],
                 $fot(empty_set)]])).

cnf(refute_0_320, plain,
    (set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) != empty_set |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(skolemFOFtoCNF_A_2, empty_set)),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)),
                      set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1))))],
              [refute_0_318, refute_0_319])).

cnf(refute_0_321, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     set_union2(skolemFOFtoCNF_A_2, empty_set)),
    inference(resolve,
              [$cnf($equal(set_difference(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1), empty_set))],
              [refute_0_309, refute_0_320])).

cnf(refute_0_322, plain,
    (set_union2(skolemFOFtoCNF_A_2, empty_set) != skolemFOFtoCNF_A_2 |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(skolemFOFtoCNF_A_2, empty_set) |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     skolemFOFtoCNF_A_2),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(skolemFOFtoCNF_A_2,
                            set_difference(skolemFOFtoCNF_A_2,
                              skolemFOFtoCNF_B_1)))),
                bind(Y, $fot(set_union2(skolemFOFtoCNF_A_2, empty_set))),
                bind(Z, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_323, plain,
    (set_union2(skolemFOFtoCNF_A_2, empty_set) != skolemFOFtoCNF_A_2 |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     skolemFOFtoCNF_A_2),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)),
                      set_union2(skolemFOFtoCNF_A_2, empty_set)))],
              [refute_0_321, refute_0_322])).

cnf(refute_0_324, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) =
     skolemFOFtoCNF_A_2),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2, empty_set),
                      skolemFOFtoCNF_A_2))],
              [refute_0_317, refute_0_323])).

cnf(refute_0_325, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     skolemFOFtoCNF_A_2 |
     skolemFOFtoCNF_A_2 =
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                          set_difference(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)),
                        set_union2(empty_set,
                          set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_A_2)))), [0],
                 $fot(skolemFOFtoCNF_A_2)]])).

cnf(refute_0_326, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     skolemFOFtoCNF_A_2 =
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2))),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)), skolemFOFtoCNF_A_2))],
              [refute_0_324, refute_0_325])).

cnf(refute_0_327, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_195, X_87))) =
     set_intersection2(X_195, X_87)),
    inference(subst, [],
              [refute_0_280 :
               [bind(B, $fot(set_intersection2(X_195, X_87)))]])).

cnf(refute_0_328, plain,
    (set_difference(A, set_difference(A, set_difference(A, X_63))) =
     set_intersection2(A, set_difference(A, X_63))),
    inference(subst, [],
              [refute_0_9 : [bind(B, $fot(set_difference(A, X_63)))]])).

cnf(refute_0_329, plain,
    (set_difference(A, set_difference(A, X_63)) =
     set_intersection2(A, X_63)),
    inference(subst, [], [refute_0_9 : [bind(B, $fot(X_63))]])).

cnf(refute_0_330, plain,
    (set_difference(A, set_difference(A, X_63)) !=
     set_intersection2(A, X_63) |
     set_difference(A, set_difference(A, set_difference(A, X_63))) !=
     set_intersection2(A, set_difference(A, X_63)) |
     set_difference(A, set_intersection2(A, X_63)) =
     set_intersection2(A, set_difference(A, X_63))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(A,
                          set_difference(A, set_difference(A, X_63))),
                        set_intersection2(A, set_difference(A, X_63)))),
                 [0, 1], $fot(set_intersection2(A, X_63))]])).

cnf(refute_0_331, plain,
    (set_difference(A, set_difference(A, set_difference(A, X_63))) !=
     set_intersection2(A, set_difference(A, X_63)) |
     set_difference(A, set_intersection2(A, X_63)) =
     set_intersection2(A, set_difference(A, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(A, set_difference(A, X_63)),
                      set_intersection2(A, X_63)))],
              [refute_0_329, refute_0_330])).

cnf(refute_0_332, plain,
    (set_difference(A, set_intersection2(A, X_63)) =
     set_intersection2(A, set_difference(A, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(A,
                        set_difference(A, set_difference(A, X_63))),
                      set_intersection2(A, set_difference(A, X_63))))],
              [refute_0_328, refute_0_331])).

cnf(refute_0_333, plain,
    (set_difference(X_83, set_intersection2(X_83, X_84)) =
     set_intersection2(X_83, set_difference(X_83, X_84))),
    inference(subst, [],
              [refute_0_332 :
               [bind(A, $fot(X_83)), bind(X_63, $fot(X_84))]])).

cnf(refute_0_334, plain,
    (set_intersection2(X_84, X_83) = set_intersection2(X_83, X_84)),
    inference(subst, [],
              [refute_0_44 : [bind(A, $fot(X_84)), bind(B, $fot(X_83))]])).

cnf(refute_0_335, plain,
    (set_intersection2(X_84, X_83) != set_intersection2(X_83, X_84) |
     set_intersection2(X_83, X_84) = set_intersection2(X_84, X_83)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_intersection2(X_84, X_83))),
                bind(Y, $fot(set_intersection2(X_83, X_84)))]])).

cnf(refute_0_336, plain,
    (set_intersection2(X_83, X_84) = set_intersection2(X_84, X_83)),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_84, X_83),
                      set_intersection2(X_83, X_84)))],
              [refute_0_334, refute_0_335])).

cnf(refute_0_337, plain,
    (set_difference(X_83, set_intersection2(X_83, X_84)) !=
     set_intersection2(X_83, set_difference(X_83, X_84)) |
     set_intersection2(X_83, X_84) != set_intersection2(X_84, X_83) |
     set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_intersection2(X_83, set_difference(X_83, X_84))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(X_83,
                          set_intersection2(X_83, X_84)),
                        set_intersection2(X_83,
                          set_difference(X_83, X_84)))), [0, 1],
                 $fot(set_intersection2(X_84, X_83))]])).

cnf(refute_0_338, plain,
    (set_difference(X_83, set_intersection2(X_83, X_84)) !=
     set_intersection2(X_83, set_difference(X_83, X_84)) |
     set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_intersection2(X_83, set_difference(X_83, X_84))),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_83, X_84),
                      set_intersection2(X_84, X_83)))],
              [refute_0_336, refute_0_337])).

cnf(refute_0_339, plain,
    (set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_intersection2(X_83, set_difference(X_83, X_84))),
    inference(resolve,
              [$cnf($equal(set_difference(X_83,
                        set_intersection2(X_83, X_84)),
                      set_intersection2(X_83,
                        set_difference(X_83, X_84))))],
              [refute_0_333, refute_0_338])).

cnf(refute_0_340, plain,
    (set_difference(A, set_intersection2(A, X_63)) !=
     set_intersection2(A, set_difference(A, X_63)) |
     set_intersection2(A, set_difference(A, X_63)) =
     set_difference(A, set_intersection2(A, X_63))),
    inference(subst, [],
              [refute_0_17 :
               [bind(X,
                     $fot(set_difference(A, set_intersection2(A, X_63)))),
                bind(Y,
                     $fot(set_intersection2(A,
                            set_difference(A, X_63))))]])).

cnf(refute_0_341, plain,
    (set_intersection2(A, set_difference(A, X_63)) =
     set_difference(A, set_intersection2(A, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(A, set_intersection2(A, X_63)),
                      set_intersection2(A, set_difference(A, X_63))))],
              [refute_0_332, refute_0_340])).

cnf(refute_0_342, plain,
    (set_intersection2(X_83, set_difference(X_83, X_84)) =
     set_difference(X_83, set_intersection2(X_83, X_84))),
    inference(subst, [],
              [refute_0_341 :
               [bind(A, $fot(X_83)), bind(X_63, $fot(X_84))]])).

cnf(refute_0_343, plain,
    (set_difference(X_83, set_intersection2(X_84, X_83)) !=
     set_intersection2(X_83, set_difference(X_83, X_84)) |
     set_intersection2(X_83, set_difference(X_83, X_84)) !=
     set_difference(X_83, set_intersection2(X_83, X_84)) |
     set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_difference(X_83, set_intersection2(X_83, X_84))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(X_83,
                          set_intersection2(X_84, X_83)),
                        set_difference(X_83,
                          set_intersection2(X_83, X_84)))), [0],
                 $fot(set_intersection2(X_83,
                        set_difference(X_83, X_84)))]])).

cnf(refute_0_344, plain,
    (set_difference(X_83, set_intersection2(X_84, X_83)) !=
     set_intersection2(X_83, set_difference(X_83, X_84)) |
     set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_difference(X_83, set_intersection2(X_83, X_84))),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_83,
                        set_difference(X_83, X_84)),
                      set_difference(X_83,
                        set_intersection2(X_83, X_84))))],
              [refute_0_342, refute_0_343])).

cnf(refute_0_345, plain,
    (set_difference(X_83, set_intersection2(X_84, X_83)) =
     set_difference(X_83, set_intersection2(X_83, X_84))),
    inference(resolve,
              [$cnf($equal(set_difference(X_83,
                        set_intersection2(X_84, X_83)),
                      set_intersection2(X_83,
                        set_difference(X_83, X_84))))],
              [refute_0_339, refute_0_344])).

cnf(refute_0_346, plain,
    (set_difference(X_87, set_intersection2(X_195, X_87)) =
     set_difference(X_87, set_intersection2(X_87, X_195))),
    inference(subst, [],
              [refute_0_345 :
               [bind(X_83, $fot(X_87)), bind(X_84, $fot(X_195))]])).

cnf(refute_0_347, plain,
    (set_difference(X_87, set_intersection2(X_195, X_87)) !=
     set_difference(X_87, set_intersection2(X_87, X_195)) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_195, X_87))) !=
     set_intersection2(X_195, X_87) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_195, X_87)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                          set_difference(X_87,
                            set_intersection2(X_195, X_87))),
                        set_intersection2(X_195, X_87))), [0, 1],
                 $fot(set_difference(X_87,
                        set_intersection2(X_87, X_195)))]])).

cnf(refute_0_348, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_195, X_87))) !=
     set_intersection2(X_195, X_87) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_195, X_87)),
    inference(resolve,
              [$cnf($equal(set_difference(X_87,
                        set_intersection2(X_195, X_87)),
                      set_difference(X_87,
                        set_intersection2(X_87, X_195))))],
              [refute_0_346, refute_0_347])).

cnf(refute_0_349, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_195, X_87)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_195, X_87))),
                      set_intersection2(X_195, X_87)))],
              [refute_0_327, refute_0_348])).

cnf(refute_0_350, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195)) = set_intersection2(X_87, X_195)),
    inference(subst, [],
              [refute_0_297 :
               [bind(X_107, $fot(X_195)), bind(X_108, $fot(X_87))]])).

cnf(refute_0_351, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_intersection2(X_208, set_difference(X_208, X_63)))),
    inference(subst, [],
              [refute_0_305 :
               [bind(X_113, $fot(X_208)),
                bind(X_114, $fot(set_difference(X_208, X_63)))]])).

cnf(refute_0_352, plain,
    (set_difference(X_208, set_intersection2(X_208, X_63)) =
     set_intersection2(X_208, set_difference(X_208, X_63))),
    inference(subst, [], [refute_0_332 : [bind(A, $fot(X_208))]])).

cnf(refute_0_353, plain,
    (set_difference(X_208, set_intersection2(X_208, X_63)) !=
     set_intersection2(X_208, set_difference(X_208, X_63)) |
     set_intersection2(X_208, set_difference(X_208, X_63)) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(subst, [],
              [refute_0_17 :
               [bind(X,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63)))),
                bind(Y,
                     $fot(set_intersection2(X_208,
                            set_difference(X_208, X_63))))]])).

cnf(refute_0_354, plain,
    (set_intersection2(X_208, set_difference(X_208, X_63)) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(X_208,
                        set_intersection2(X_208, X_63)),
                      set_intersection2(X_208,
                        set_difference(X_208, X_63))))],
              [refute_0_352, refute_0_353])).

cnf(refute_0_355, plain,
    (set_intersection2(X_208, set_difference(X_208, X_63)) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_intersection2(X_208, set_difference(X_208, X_63))) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63)))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_208, X_63),
                          set_difference(set_difference(X_208, X_63),
                            X_208)),
                        set_union2(set_difference(set_difference(X_208,
                              X_63), X_208),
                          set_intersection2(X_208,
                            set_difference(X_208, X_63))))), [1, 1],
                 $fot(set_difference(X_208,
                        set_intersection2(X_208, X_63)))]])).

cnf(refute_0_356, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_intersection2(X_208, set_difference(X_208, X_63))) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63)))),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_208,
                        set_difference(X_208, X_63)),
                      set_difference(X_208,
                        set_intersection2(X_208, X_63))))],
              [refute_0_354, refute_0_355])).

cnf(refute_0_357, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63)))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(set_difference(X_208,
                            X_63), X_208),
                        set_intersection2(X_208,
                          set_difference(X_208, X_63)))))],
              [refute_0_351, refute_0_356])).

cnf(refute_0_358, plain,
    (set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63))) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187,
                     $fot(set_difference(set_difference(X_208, X_63),
                            X_208))),
                bind(X_188,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63))))]])).

cnf(refute_0_359, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63))) |
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63))) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_union2(set_difference(X_208, X_63),
                          set_difference(set_difference(X_208, X_63),
                            X_208)),
                        set_union2(set_difference(X_208,
                            set_intersection2(X_208, X_63)),
                          set_difference(set_difference(X_208, X_63),
                            X_208)))), [0],
                 $fot(set_union2(set_difference(set_difference(X_208,
                            X_63), X_208),
                        set_difference(X_208,
                          set_intersection2(X_208, X_63))))]])).

cnf(refute_0_360, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(set_difference(X_208, X_63), X_208),
       set_difference(X_208, set_intersection2(X_208, X_63))) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(set_difference(X_208,
                            X_63), X_208),
                        set_difference(X_208,
                          set_intersection2(X_208, X_63))),
                      set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208))))], [refute_0_358, refute_0_359])).

cnf(refute_0_361, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(set_difference(X_208,
                            X_63), X_208),
                        set_difference(X_208,
                          set_intersection2(X_208, X_63)))))],
              [refute_0_357, refute_0_360])).

cnf(refute_0_362, plain,
    (set_union2(empty_set, set_difference(X_208, X_63)) =
     set_difference(X_208, X_63)),
    inference(subst, [],
              [refute_0_36 :
               [bind(A, $fot(set_difference(X_208, X_63)))]])).

cnf(refute_0_363, plain,
    (set_union2(set_difference(X_208, X_63), empty_set) =
     set_union2(empty_set, set_difference(X_208, X_63))),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187, $fot(set_difference(X_208, X_63))),
                bind(X_188, $fot(empty_set))]])).

cnf(refute_0_364, plain,
    (set_union2(empty_set, set_difference(X_208, X_63)) !=
     set_difference(X_208, X_63) |
     set_union2(set_difference(X_208, X_63), empty_set) !=
     set_union2(empty_set, set_difference(X_208, X_63)) |
     set_union2(set_difference(X_208, X_63), empty_set) =
     set_difference(X_208, X_63)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(X_208, X_63),
                            empty_set))),
                bind(Y,
                     $fot(set_union2(empty_set,
                            set_difference(X_208, X_63)))),
                bind(Z, $fot(set_difference(X_208, X_63)))]])).

cnf(refute_0_365, plain,
    (set_union2(empty_set, set_difference(X_208, X_63)) !=
     set_difference(X_208, X_63) |
     set_union2(set_difference(X_208, X_63), empty_set) =
     set_difference(X_208, X_63)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        empty_set),
                      set_union2(empty_set,
                        set_difference(X_208, X_63))))],
              [refute_0_363, refute_0_364])).

cnf(refute_0_366, plain,
    (set_union2(set_difference(X_208, X_63), empty_set) =
     set_difference(X_208, X_63)),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set,
                        set_difference(X_208, X_63)),
                      set_difference(X_208, X_63)))],
              [refute_0_362, refute_0_365])).

cnf(refute_0_367, plain,
    (set_difference(set_difference(X_208, X_63), X_208) = empty_set),
    inference(subst, [],
              [refute_0_206 :
               [bind(B, $fot(X_63)), bind(X_236, $fot(X_208))]])).

cnf(refute_0_368, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208))),
    introduced(tautology,
               [refl,
                [$fot(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)))]])).

cnf(refute_0_369, plain,
    (set_difference(set_difference(X_208, X_63), X_208) != empty_set |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, X_63), empty_set)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_208, X_63),
                          set_difference(set_difference(X_208, X_63),
                            X_208)),
                        set_union2(set_difference(X_208, X_63),
                          set_difference(set_difference(X_208, X_63),
                            X_208)))), [1, 1], $fot(empty_set)]])).

cnf(refute_0_370, plain,
    (set_difference(set_difference(X_208, X_63), X_208) != empty_set |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, X_63), empty_set)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208))))], [refute_0_368, refute_0_369])).

cnf(refute_0_371, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, X_63), empty_set)),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_208, X_63),
                        X_208), empty_set))],
              [refute_0_367, refute_0_370])).

cnf(refute_0_372, plain,
    (set_union2(set_difference(X_208, X_63), empty_set) !=
     set_difference(X_208, X_63) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, X_63), empty_set) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, X_63)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(X_208, X_63),
                            set_difference(set_difference(X_208, X_63),
                              X_208)))),
                bind(Y,
                     $fot(set_union2(set_difference(X_208, X_63),
                            empty_set))),
                bind(Z, $fot(set_difference(X_208, X_63)))]])).

cnf(refute_0_373, plain,
    (set_union2(set_difference(X_208, X_63), empty_set) !=
     set_difference(X_208, X_63) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, X_63)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(X_208, X_63),
                        empty_set)))], [refute_0_371, refute_0_372])).

cnf(refute_0_374, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, X_63)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        empty_set), set_difference(X_208, X_63)))],
              [refute_0_366, refute_0_373])).

cnf(refute_0_375, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_difference(X_208, X_63) |
     set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_difference(X_208, X_63) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_208, X_63),
                          set_difference(set_difference(X_208, X_63),
                            X_208)),
                        set_union2(set_difference(X_208,
                            set_intersection2(X_208, X_63)),
                          set_difference(set_difference(X_208, X_63),
                            X_208)))), [0],
                 $fot(set_difference(X_208, X_63))]])).

cnf(refute_0_376, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_difference(X_208, X_63) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)), set_difference(X_208, X_63)))],
              [refute_0_374, refute_0_375])).

cnf(refute_0_377, plain,
    (set_union2(empty_set,
       set_difference(X_208, set_intersection2(X_208, X_63))) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(subst, [],
              [refute_0_36 :
               [bind(A,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63))))]])).

cnf(refute_0_378, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) =
     set_union2(empty_set,
       set_difference(X_208, set_intersection2(X_208, X_63)))),
    inference(subst, [],
              [refute_0_193 :
               [bind(X_187,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63)))),
                bind(X_188, $fot(empty_set))]])).

cnf(refute_0_379, plain,
    (set_union2(empty_set,
       set_difference(X_208, set_intersection2(X_208, X_63))) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) !=
     set_union2(empty_set,
       set_difference(X_208, set_intersection2(X_208, X_63))) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) = set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(X_208,
                              set_intersection2(X_208, X_63)),
                            empty_set))),
                bind(Y,
                     $fot(set_union2(empty_set,
                            set_difference(X_208,
                              set_intersection2(X_208, X_63))))),
                bind(Z,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63))))]])).

cnf(refute_0_380, plain,
    (set_union2(empty_set,
       set_difference(X_208, set_intersection2(X_208, X_63))) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) = set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)), empty_set),
                      set_union2(empty_set,
                        set_difference(X_208,
                          set_intersection2(X_208, X_63)))))],
              [refute_0_378, refute_0_379])).

cnf(refute_0_381, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) = set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set,
                        set_difference(X_208,
                          set_intersection2(X_208, X_63))),
                      set_difference(X_208,
                        set_intersection2(X_208, X_63))))],
              [refute_0_377, refute_0_380])).

cnf(refute_0_382, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208))),
    introduced(tautology,
               [refl,
                [$fot(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208)))]])).

cnf(refute_0_383, plain,
    (set_difference(set_difference(X_208, X_63), X_208) != empty_set |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_union2(set_difference(X_208,
                            set_intersection2(X_208, X_63)),
                          set_difference(set_difference(X_208, X_63),
                            X_208)),
                        set_union2(set_difference(X_208,
                            set_intersection2(X_208, X_63)),
                          set_difference(set_difference(X_208, X_63),
                            X_208)))), [1, 1], $fot(empty_set)]])).

cnf(refute_0_384, plain,
    (set_difference(set_difference(X_208, X_63), X_208) != empty_set |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set)),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208))))], [refute_0_382, refute_0_383])).

cnf(refute_0_385, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set)),
    inference(resolve,
              [$cnf($equal(set_difference(set_difference(X_208, X_63),
                        X_208), empty_set))],
              [refute_0_367, refute_0_384])).

cnf(refute_0_386, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(set_difference(X_208,
                              set_intersection2(X_208, X_63)),
                            set_difference(set_difference(X_208, X_63),
                              X_208)))),
                bind(Y,
                     $fot(set_union2(set_difference(X_208,
                              set_intersection2(X_208, X_63)),
                            empty_set))),
                bind(Z,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63))))]])).

cnf(refute_0_387, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       empty_set) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)), empty_set)))],
              [refute_0_385, refute_0_386])).

cnf(refute_0_388, plain,
    (set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)), empty_set),
                      set_difference(X_208,
                        set_intersection2(X_208, X_63))))],
              [refute_0_381, refute_0_387])).

cnf(refute_0_389, plain,
    (set_difference(X_208, X_63) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_difference(X_208, X_63) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    introduced(tautology,
               [equality,
                [$cnf(~ $equal(set_difference(X_208, X_63),
                        set_difference(X_208,
                          set_intersection2(X_208, X_63)))), [0],
                 $fot(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208)))]])).

cnf(refute_0_390, plain,
    (set_difference(X_208, X_63) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_difference(X_208, X_63) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_difference(X_208,
                        set_intersection2(X_208, X_63))))],
              [refute_0_388, refute_0_389])).

cnf(refute_0_391, plain,
    (set_union2(set_difference(X_208, X_63),
       set_difference(set_difference(X_208, X_63), X_208)) !=
     set_union2(set_difference(X_208, set_intersection2(X_208, X_63)),
       set_difference(set_difference(X_208, X_63), X_208)) |
     set_difference(X_208, X_63) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_difference(X_208, X_63),
                      set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208))))], [refute_0_376, refute_0_390])).

cnf(refute_0_392, plain,
    (set_difference(X_208, X_63) =
     set_difference(X_208, set_intersection2(X_208, X_63))),
    inference(resolve,
              [$cnf($equal(set_union2(set_difference(X_208, X_63),
                        set_difference(set_difference(X_208, X_63),
                          X_208)),
                      set_union2(set_difference(X_208,
                          set_intersection2(X_208, X_63)),
                        set_difference(set_difference(X_208, X_63),
                          X_208))))], [refute_0_361, refute_0_391])).

cnf(refute_0_393, plain,
    (set_difference(X_208, X_63) !=
     set_difference(X_208, set_intersection2(X_208, X_63)) |
     set_difference(X_208, set_intersection2(X_208, X_63)) =
     set_difference(X_208, X_63)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_difference(X_208, X_63))),
                bind(Y,
                     $fot(set_difference(X_208,
                            set_intersection2(X_208, X_63))))]])).

cnf(refute_0_394, plain,
    (set_difference(X_208, set_intersection2(X_208, X_63)) =
     set_difference(X_208, X_63)),
    inference(resolve,
              [$cnf($equal(set_difference(X_208, X_63),
                      set_difference(X_208,
                        set_intersection2(X_208, X_63))))],
              [refute_0_392, refute_0_393])).

cnf(refute_0_395, plain,
    (set_difference(X_87, set_intersection2(X_87, X_195)) =
     set_difference(X_87, X_195)),
    inference(subst, [],
              [refute_0_394 :
               [bind(X_208, $fot(X_87)), bind(X_63, $fot(X_195))]])).

cnf(refute_0_396, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195)))),
    introduced(tautology,
               [refl,
                [$fot(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195))))]])).

cnf(refute_0_397, plain,
    (set_difference(X_87, set_intersection2(X_87, X_195)) !=
     set_difference(X_87, X_195) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) !=
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195))),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                          set_difference(X_87,
                            set_intersection2(X_87, X_195))),
                        set_difference(set_intersection2(X_195, X_87),
                          set_difference(X_87,
                            set_intersection2(X_87, X_195))))), [1, 1],
                 $fot(set_difference(X_87, X_195))]])).

cnf(refute_0_398, plain,
    (set_difference(X_87, set_intersection2(X_87, X_195)) !=
     set_difference(X_87, X_195) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195))),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195))),
                      set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195)))))],
              [refute_0_396, refute_0_397])).

cnf(refute_0_399, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195))),
    inference(resolve,
              [$cnf($equal(set_difference(X_87,
                        set_intersection2(X_87, X_195)),
                      set_difference(X_87, X_195)))],
              [refute_0_395, refute_0_398])).

cnf(refute_0_400, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195)) != set_intersection2(X_87, X_195) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) !=
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195)) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_87, X_195)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_difference(set_intersection2(X_195, X_87),
                            set_difference(X_87,
                              set_intersection2(X_87, X_195))))),
                bind(Y,
                     $fot(set_difference(set_intersection2(X_195, X_87),
                            set_difference(X_87, X_195)))),
                bind(Z, $fot(set_intersection2(X_87, X_195)))]])).

cnf(refute_0_401, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, X_195)) != set_intersection2(X_87, X_195) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_87, X_195)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195))),
                      set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87, X_195))))],
              [refute_0_399, refute_0_400])).

cnf(refute_0_402, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) =
     set_intersection2(X_87, X_195)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87, X_195)),
                      set_intersection2(X_87, X_195)))],
              [refute_0_350, refute_0_401])).

cnf(refute_0_403, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) !=
     set_intersection2(X_195, X_87) |
     set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) !=
     set_intersection2(X_87, X_195) |
     set_intersection2(X_87, X_195) = set_intersection2(X_195, X_87)),
    introduced(tautology,
               [equality,
                [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                          set_difference(X_87,
                            set_intersection2(X_87, X_195))),
                        set_intersection2(X_195, X_87))), [0],
                 $fot(set_intersection2(X_87, X_195))]])).

cnf(refute_0_404, plain,
    (set_difference(set_intersection2(X_195, X_87),
       set_difference(X_87, set_intersection2(X_87, X_195))) !=
     set_intersection2(X_195, X_87) |
     set_intersection2(X_87, X_195) = set_intersection2(X_195, X_87)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195))),
                      set_intersection2(X_87, X_195)))],
              [refute_0_402, refute_0_403])).

cnf(refute_0_405, plain,
    (set_intersection2(X_87, X_195) = set_intersection2(X_195, X_87)),
    inference(resolve,
              [$cnf($equal(set_difference(set_intersection2(X_195, X_87),
                        set_difference(X_87,
                          set_intersection2(X_87, X_195))),
                      set_intersection2(X_195, X_87)))],
              [refute_0_349, refute_0_404])).

cnf(refute_0_406, plain,
    (set_intersection2(X_87, X_195) != set_intersection2(X_195, X_87) |
     set_intersection2(X_195, X_87) = set_intersection2(X_87, X_195)),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(set_intersection2(X_87, X_195))),
                bind(Y, $fot(set_intersection2(X_195, X_87)))]])).

cnf(refute_0_407, plain,
    (set_intersection2(X_195, X_87) = set_intersection2(X_87, X_195)),
    inference(resolve,
              [$cnf($equal(set_intersection2(X_87, X_195),
                      set_intersection2(X_195, X_87)))],
              [refute_0_405, refute_0_406])).

cnf(refute_0_408, plain,
    (set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2) =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(subst, [],
              [refute_0_407 :
               [bind(X_195, $fot(skolemFOFtoCNF_B_1)),
                bind(X_87, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_409, plain,
    (set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) =
     set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)),
    inference(subst, [],
              [refute_0_36 :
               [bind(A,
                     $fot(set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_A_2)))]])).

cnf(refute_0_410, plain,
    (set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2) !=
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) |
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) !=
     set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2) |
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(subst, [],
              [refute_0_71 :
               [bind(X,
                     $fot(set_union2(empty_set,
                            set_intersection2(skolemFOFtoCNF_B_1,
                              skolemFOFtoCNF_A_2)))),
                bind(Y,
                     $fot(set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_A_2))),
                bind(Z,
                     $fot(set_intersection2(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)))]])).

cnf(refute_0_411, plain,
    (set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2) !=
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) |
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set,
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_A_2)),
                      set_intersection2(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_A_2)))],
              [refute_0_409, refute_0_410])).

cnf(refute_0_412, plain,
    (set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf($equal(set_intersection2(skolemFOFtoCNF_B_1,
                        skolemFOFtoCNF_A_2),
                      set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1)))],
              [refute_0_408, refute_0_411])).

cnf(refute_0_413, plain,
    (set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) !=
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) |
     skolemFOFtoCNF_A_2 !=
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     skolemFOFtoCNF_A_2 =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    introduced(tautology,
               [equality,
                [$cnf($equal(skolemFOFtoCNF_A_2,
                        set_union2(empty_set,
                          set_intersection2(skolemFOFtoCNF_B_1,
                            skolemFOFtoCNF_A_2)))), [1],
                 $fot(set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1))]])).

cnf(refute_0_414, plain,
    (skolemFOFtoCNF_A_2 !=
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     skolemFOFtoCNF_A_2 =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf($equal(set_union2(empty_set,
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_A_2)),
                      set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1)))],
              [refute_0_412, refute_0_413])).

cnf(refute_0_415, plain,
    (set_union2(skolemFOFtoCNF_A_2,
       set_difference(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)) !=
     set_union2(empty_set,
       set_intersection2(skolemFOFtoCNF_B_1, skolemFOFtoCNF_A_2)) |
     skolemFOFtoCNF_A_2 =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_A_2,
                      set_union2(empty_set,
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_A_2))))],
              [refute_0_326, refute_0_414])).

cnf(refute_0_416, plain,
    (skolemFOFtoCNF_A_2 =
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf($equal(set_union2(skolemFOFtoCNF_A_2,
                        set_difference(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1)),
                      set_union2(empty_set,
                        set_intersection2(skolemFOFtoCNF_B_1,
                          skolemFOFtoCNF_A_2))))],
              [refute_0_312, refute_0_415])).

cnf(refute_0_417, plain,
    (skolemFOFtoCNF_A_2 !=
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) |
     set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) =
     skolemFOFtoCNF_A_2),
    inference(subst, [],
              [refute_0_17 :
               [bind(X, $fot(skolemFOFtoCNF_A_2)),
                bind(Y,
                     $fot(set_intersection2(skolemFOFtoCNF_A_2,
                            skolemFOFtoCNF_B_1)))]])).

cnf(refute_0_418, plain,
    (set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) =
     skolemFOFtoCNF_A_2),
    inference(resolve,
              [$cnf($equal(skolemFOFtoCNF_A_2,
                      set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1)))],
              [refute_0_416, refute_0_417])).

cnf(refute_0_419, plain,
    (set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1) !=
     skolemFOFtoCNF_A_2 | ~ in(X_336, skolemFOFtoCNF_A_2) |
     in(X_336, set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1))),
    introduced(tautology,
               [equality,
                [$cnf(~ in(X_336,
                        set_intersection2(skolemFOFtoCNF_A_2,
                          skolemFOFtoCNF_B_1))), [1],
                 $fot(skolemFOFtoCNF_A_2)]])).

cnf(refute_0_420, plain,
    (~ in(X_336, skolemFOFtoCNF_A_2) |
     in(X_336, set_intersection2(skolemFOFtoCNF_A_2, skolemFOFtoCNF_B_1))),
    inference(resolve,
              [$cnf($equal(set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1), skolemFOFtoCNF_A_2))],
              [refute_0_418, refute_0_419])).

cnf(refute_0_421, plain,
    (~ in(X_336, skolemFOFtoCNF_A_2) | in(X_336, skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf(in(X_336,
                      set_intersection2(skolemFOFtoCNF_A_2,
                        skolemFOFtoCNF_B_1)))],
              [refute_0_420, refute_0_241])).

cnf(refute_0_422, plain,
    (~
       in(skolemFOFtoCNF_C_3(X_341, skolemFOFtoCNF_A_2),
         skolemFOFtoCNF_A_2) |
     in(skolemFOFtoCNF_C_3(X_341, skolemFOFtoCNF_A_2),
       skolemFOFtoCNF_B_1)),
    inference(subst, [],
              [refute_0_421 :
               [bind(X_336,
                     $fot(skolemFOFtoCNF_C_3(X_341,
                            skolemFOFtoCNF_A_2)))]])).

cnf(refute_0_423, plain,
    (disjoint(X_334, X_335) |
     in(skolemFOFtoCNF_C_3(X_334, X_335),
       set_intersection2(X_334, X_335))),
    inference(subst, [],
              [refute_0_227 :
               [bind(A, $fot(X_334)), bind(B, $fot(X_335))]])).

cnf(refute_0_424, plain,
    (~
       in(skolemFOFtoCNF_C_3(X_334, X_335),
         set_intersection2(X_334, X_335)) |
     in(skolemFOFtoCNF_C_3(X_334, X_335), X_335)),
    inference(subst, [],
              [refute_0_240 :
               [bind(A, $fot(X_334)), bind(B, $fot(X_335)),
                bind(D, $fot(skolemFOFtoCNF_C_3(X_334, X_335)))]])).

cnf(refute_0_425, plain,
    (disjoint(X_334, X_335) | in(skolemFOFtoCNF_C_3(X_334, X_335), X_335)),
    inference(resolve,
              [$cnf(in(skolemFOFtoCNF_C_3(X_334, X_335),
                      set_intersection2(X_334, X_335)))],
              [refute_0_423, refute_0_424])).

cnf(refute_0_426, plain,
    (disjoint(X_341, skolemFOFtoCNF_A_2) |
     in(skolemFOFtoCNF_C_3(X_341, skolemFOFtoCNF_A_2),
       skolemFOFtoCNF_A_2)),
    inference(subst, [],
              [refute_0_425 :
               [bind(X_334, $fot(X_341)),
                bind(X_335, $fot(skolemFOFtoCNF_A_2))]])).

cnf(refute_0_427, plain,
    (disjoint(X_341, skolemFOFtoCNF_A_2) |
     in(skolemFOFtoCNF_C_3(X_341, skolemFOFtoCNF_A_2),
       skolemFOFtoCNF_B_1)),
    inference(resolve,
              [$cnf(in(skolemFOFtoCNF_C_3(X_341, skolemFOFtoCNF_A_2),
                      skolemFOFtoCNF_A_2))],
              [refute_0_426, refute_0_422])).

cnf(refute_0_428, plain,
    (disjoint(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2) |
     in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2),
       skolemFOFtoCNF_B_1)),
    inference(subst, [],
              [refute_0_427 : [bind(X_341, $fot(skolemFOFtoCNF_C_4))]])).

cnf(refute_0_429, plain,
    (disjoint(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2)),
    inference(resolve,
              [$cnf(in(skolemFOFtoCNF_C_3(skolemFOFtoCNF_C_4,
                        skolemFOFtoCNF_A_2), skolemFOFtoCNF_B_1))],
              [refute_0_428, refute_0_237])).

cnf(refute_0_430, plain,
    (disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4)),
    inference(resolve,
              [$cnf(disjoint(skolemFOFtoCNF_C_4, skolemFOFtoCNF_A_2))],
              [refute_0_429, refute_0_1])).

cnf(refute_0_431, plain,
    (~ disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4)),
    inference(canonicalize, [], [normalize_0_45])).

cnf(refute_0_432, plain, ($false),
    inference(resolve,
              [$cnf(disjoint(skolemFOFtoCNF_A_2, skolemFOFtoCNF_C_4))],
              [refute_0_430, refute_0_431])).
% SZS output end CNFRefutation for data/problems/all/SEU140+2.tptp

